Before you turn this problem in, make sure everything runs as expected. 
  1. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and 
  2. Then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name below. 

---

# CSCI 3155: Assignment 2

Topics covered: recursion and inductive definitions.

__YOUR NAME HERE__

Note: Please first run the `TEST HELPER` cell that defines the `passed` function below. Failing to run this cell will make it hard for you to check your work.

```scala
// TEST HELPER
def passed(points: Int) {
    require(points >=0)
    if (points == 1) print(s"\n*** Tests Passed (1 point) ***\n")
    else print(s"\n*** Tests Passed ($points points) ***\n")
}
```

Output:
```
defined [32mfunction[39m [36mpassed[39m
```

## Problem 1 (10 points total)

### A (5 points) Recursively compute binary representation

Given a positive integer $n$, we wish to compute its binary representation as a string. 

  - Use of var and loops forbidden. 
  - Inbuilt functions that directly perform conversion are also forbidden.

__Example:__ if input is 5, output should be "101".

Use the following scheme:
  - If number is 0 or 1, we can convert them into binary string "0" or "1", respectively.
  - If input number `n` is odd, recursively convert `(n-1)/2` into binary and append a `1` to the end of the string returned.
  - else  (input number `n` is even) convert `n/2` into binary and append `0` to end of the string returned.


<!-- grade_id: answer_1a, points: 0, grade: false, solution: true -->
```scala
def convertToBinary(n: Int): String = {
    require( n >= 0) // If this is violated, you will get a "requirement failed message"
    // YOUR CODE HERE
    if(n==0 | n==1)
    {
        return n.toString
    }
    else
    {
        if(n%2==1)
        {
            return convertToBinary( (n-1)/2)+1
        }
        return convertToBinary( n/2)+0
    }
}
```

Output:
```
defined [32mfunction[39m [36mconvertToBinary[39m
```

<!-- grade_id: tests_1a, points: 5, grade: true, solution: false -->
```scala
//BEGIN TEST
assert(convertToBinary(1) == "1", "(1)decimal must be 1 binary")
assert(convertToBinary(5) == "101", "5 must be 101 binary")
assert(convertToBinary(11) == "1011", "(11) must be 1011")
assert(convertToBinary(4100) == "1000000000100", "(4100) must be 1000000000100")
assert(convertToBinary(30108100) == "1110010110110100111000100", "(30108100) must be 1110010110110100111000100")
passed(5)
//END TEST
```

Output:
```
*** Tests Passed (5 points) ***
```

### B (5 points) _Tail_ Recursive `convertToBinary`
Now, write a tail recursive version of the `convertToBinary` function, called `convertToBinaryTail`. 

You should implement the helper function as a tail recursive function. 

Use of var and loops forbidden. Inbuilt functions that directly perform conversion are also forbidden.

<!-- grade_id: answer_1b, points: 0, grade: false, solution: true -->
```scala
import scala.annotation.tailrec

@tailrec
final def convertToBinaryHelper(n: Int, stringSoFar: String): String = {
    require(n >= 1)
    // YOUR CODE HERE
    if(n==1){ return "1"+stringSoFar}
    else{
      // println("ll")
     if(n%2==1)
         { 
             convertToBinaryHelper( (n-1)/2,1+stringSoFar)
           
         }
        else
        { 
            convertToBinaryHelper( n/2,0+stringSoFar)
        }
       
    
        
    }

    
}

def convertToBinaryTail(n: Int): String = { 
    // YOUR CODE HERE
            return convertToBinaryHelper( n,"")
}

```

Output:
```
[32mimport [39m[36mscala.annotation.tailrec

[39m
defined [32mfunction[39m [36mconvertToBinaryHelper[39m
defined [32mfunction[39m [36mconvertToBinaryTail[39m
```

<!-- grade_id: tests_1b, points: 5, grade: true, solution: false -->
```scala
// BEGIN TEST
assert(convertToBinaryTail(1) == "1", "(1)decimal must be 1 binary")
assert(convertToBinaryTail(5) == "101", "5 must be 101 binary")
assert(convertToBinaryTail(11) == "1011", "(11) must be 1011")
assert(convertToBinaryTail(4100) == "1000000000100", "(4100) must be 1000000000100")
assert(convertToBinaryTail(30108100) == "1110010110110100111000100", "(30108100) must be 1110010110110100111000100")
passed(5)
//END TEST
```

Output:
```
*** Tests Passed (5 points) ***
```

## Problem 2 (18 points)

## Part A
Convert the iterative function below provided as a loop into a tail recursive version.


```scala
def seriesSummationUsingLoop(a: Int, b: Int, func: Int => Int): Int = {
    var acc = 0
    var i = a
    while (i <= b){
        acc = acc + func(i)
        i = i + 1
    }
    return acc
}
```

Implement a tail recursive version using an accumulator argument to the function and a counter argument as shown below.
Use of var/loops is forbidden.

<!-- grade_id: answer_2, points: 0, grade: false, solution: true -->
```scala
import scala.annotation.tailrec

@tailrec
def tailRecursiveSummationHelper(b: Int, func: Int => Int, i: Int, acc: Int): Int = {
    assert(i <= b)
    // YOUR CODE HERE
    if(i==b){return acc+func(i)}
    else
    { 
       return  tailRecursiveSummationHelper(b,func,i+1,acc+func(i)) 
       
    }
  
}

def tailRecursiveSummation(a: Int, b: Int, func: Int => Int) = {
    tailRecursiveSummationHelper(b, func, a, 0)
}
```

Output:
```
[32mimport [39m[36mscala.annotation.tailrec

[39m
defined [32mfunction[39m [36mtailRecursiveSummationHelper[39m
defined [32mfunction[39m [36mtailRecursiveSummation[39m
```

<!-- grade_id: tests_2, points: 8, grade: true, solution: false -->
```scala
//BEGIN TEST
val square = (i: Int) => i*i
val ident = (i:Int) => i
assert(tailRecursiveSummation(1, 10, ident) == 55, "summation of 1 to 10 is 55")
assert(tailRecursiveSummation(1, 4, square) == 30, "summation of 1^2+2^2+3^2+4^2 is 30")
assert(tailRecursiveSummation(4, 4, square) == 16, "summation of 4^2 is 16")
assert(tailRecursiveSummation(-1, 1, ident) == 0, "summation of -1+0+1 is 0")

passed(8)
//END TEST
```

Output:
```
*** Tests Passed (8 points) ***
[36msquare[39m: [32mInt[39m => [32mInt[39m = ammonite.$sess.cmd123$Helper$$Lambda$3001/0x0000000840cf8840@27ede44f
[36mident[39m: [32mInt[39m => [32mInt[39m = ammonite.$sess.cmd123$Helper$$Lambda$3002/0x0000000840cf9840@6bec1f77
```

## Part B
Convert the iterative function below provided as a loop into a tail recursive version.

The function checks if `pat` occurs as a substring in string `str`.
But rather than use a nested loop, it uses a single while loop.


```scala
def findPatternInString(str: String, pat: String): Boolean = {
    val n1 = str.size
    val n2 = pat.size
    var i = 0
    var j = 0
    while (i < n1 - n2 + 1){
        if ( j < n2 && str(i+j) == pat(j)){
            j = j + 1
        } else if (j == n2) {
            return true
        } else { 
            i = i + 1
            j = 0
        }
    }
    return false
}
```

Output:
```
defined [32mfunction[39m [36mfindPatternInString[39m
```

```scala
findPatternInString("hello", "ll")
findPatternInString("hello", "ellx")
findPatternInString("hello", "ello")
findPatternInString("hello", "hellx")
findPatternInString("hello", "helloo")
findPatternInString("hello", "hello")
findPatternInString("hello", "")
findPatternInString("heheeheeeheeee", "heeee")
findPatternInString("hello", "helo")
```

Output:
```
[36mres168_0[39m: [32mBoolean[39m = true
[36mres168_1[39m: [32mBoolean[39m = false
[36mres168_2[39m: [32mBoolean[39m = true
[36mres168_3[39m: [32mBoolean[39m = false
[36mres168_4[39m: [32mBoolean[39m = false
[36mres168_5[39m: [32mBoolean[39m = true
[36mres168_6[39m: [32mBoolean[39m = true
[36mres168_7[39m: [32mBoolean[39m = true
[36mres168_8[39m: [32mBoolean[39m = false
```

Complete the following tail recursive version by filling in missing parts so that its behavior is the same as the iterative version above.

Restrictions:
 - Must be tail recursive.
 - No vars/loops allowed.
 
Try not to use return but we will not penalize you if you use return.

<!-- grade_id: cell-df05298b72cd465d, points: 0, grade: false, solution: true -->
```scala
def findPatternInStringRec(str: String, pat: String, i: Int= 0, j: Int =0): Boolean = {
    val n1 = str.size
    val n2 = pat.size
    if (i < n1 - n2 + 1){
        if ( j < n2 && str(i+j) == pat(j)){
            // YOUR CODE HERE
            findPatternInStringRec(str,pat,i,j+1)
        } else if (j == n2) {
            // YOUR CODE HERE
            true
        } else { 
            // YOUR CODE HERE
            findPatternInStringRec(str,pat,i+1,0)
        }
    } else { 
        false
    }
}
```

Output:
```
defined [32mfunction[39m [36mfindPatternInStringRec[39m
```

<!-- grade_id: cell-4d5a2b813f6cbcc3, points: 10, grade: true, solution: false -->
```scala
assert(findPatternInString("hello", "ll"), "Test 1 failed")
assert(!findPatternInString("hello", "ellx"), "Test 2 failed")
assert(findPatternInString("hello", "ello"),  "Test 3 failed")
assert(!findPatternInString("hello", "hellx"), "Test 4 failed")
assert(!findPatternInString("hello", "helloo"), "Test 5 failed")
assert(findPatternInString("hello", "hello"), "Test 6 failed")
assert(findPatternInString("hello", ""), "Test 7 failed")
assert(findPatternInString("heheeheeeheeee", "heeee"), "Test 8 failed")
assert(!findPatternInString("hello", "helo"), "Test 9 failed")
passed(10)
```

Output:
```
*** Tests Passed (10 points) ***
```

## Problem 3 (22 points)

Convert the following inductive definition for regular expressions into a grammar first and then into a set of scala classes.

A regular expression is defined inductively as follows:
- Any string s is an "atomic" regular expression.
- If $r_1$, $r_2$ are regular expressions then so are
  - The concatenation $r_1 ; r_2$, 
  - The disjunction $r_1 | r_2$, and 
  - The conjunction $r_1 \& r_2$.
- If $r$ is a regular expression, then its Kleene star $r^*$ is also a regular expression.

Use the constructor symbols:
  - $\text{Atom}(s)$ to denote an atomic regular expression,
  - $\text{Concatenation}(r_1, r_2)$ for the ";" operator, 
  - $\text{OrOperator}(r_1, r_2)$ for the "|" operator, 
  - $\text{AndOperator}(r_1, r_2)$ for the "&" operator and 
  - $\text{KleeneStar}(r)$ for the Kleene-star operator. 
  
You may use the nonterminal $\textbf{string}$ without definition to denote a string of characters.




### A (7 point)
Write the grammar using constructor symbols for the inductive definition above. __Tip:__ you can examine the notebooks with inductive definitions to see how we typeset the grammar. There are no tests for this because it will be manualy graded.

<!-- grade_id: answer_3a, points: 7, grade: true, solution: true -->
YOUR ANSWER HERE
 $$\begin{array}
  \textbf{Regex} & \rightarrow&  Atom(\textbf{string}) \\ 
   \textbf{Regex} & \rightarrow&  Concatenation(\textbf{Regex},\textbf{Regex}) \\ 
    \textbf{Regex} & \rightarrow&  OrOperator(\textbf{Regex},\textbf{Regex}) \\\
     \textbf{Regex} & \rightarrow&   KleeneStar(\textbf{Regex}) \\ 
      \textbf{Regex} & \rightarrow&  AndOperator(\textbf{Regex},\textbf{Regex}) \\\ 
                  \end{array}$$ 

### B (7 points)
Define the structure as a set of case class in scala.


<!-- grade_id: answer_3b, points: 0, grade: false, solution: true -->
```scala
sealed trait Regex
// Use constructors: Atom, Concatenation, OrOperator, AndOperator and KleeneStar
// YOUR CODE HERE
case class Atom(s:String) extends Regex
case class Concatenation(v1:Regex,v2:Regex) extends Regex
case class OrOperator(v1:Regex,v2:Regex) extends Regex
case class KleeneStar(v3:Regex) extends Regex
case class AndOperator(v1:Regex, v4:Regex)extends Regex
```

Output:
```
defined [32mtrait[39m [36mRegex[39m
defined [32mclass[39m [36mAtom[39m
defined [32mclass[39m [36mConcatenation[39m
defined [32mclass[39m [36mOrOperator[39m
defined [32mclass[39m [36mKleeneStar[39m
defined [32mclass[39m [36mAndOperator[39m
```

<!-- grade_id: tests_3b, points: 7, grade: true, solution: false -->
```scala
//BEGIN TEST
val v1 = Atom("Hello")
val v2 = Concatenation(v1, v1)
val v3 = OrOperator(v1, v2)
val v4 = KleeneStar(v3)
val v5 = AndOperator(v1, v4)
passed(10)
//END TEST
```

Output:
```
*** Tests Passed (10 points) ***
[36mv1[39m: [32mAtom[39m = [33mAtom[39m([32m"Hello"[39m)
[36mv2[39m: [32mConcatenation[39m = [33mConcatenation[39m([33mAtom[39m([32m"Hello"[39m), [33mAtom[39m([32m"Hello"[39m))
[36mv3[39m: [32mOrOperator[39m = [33mOrOperator[39m(
  [33mAtom[39m([32m"Hello"[39m),
  [33mConcatenation[39m([33mAtom[39m([32m"Hello"[39m), [33mAtom[39m([32m"Hello"[39m))
)
[36mv4[39m: [32mKleeneStar[39m = [33mKleeneStar[39m(
  [33mOrOperator[39m([33mAtom[39m([32m"Hello"[39m), [33mConcatenation[39m([33mAtom[39m([32m"Hello"[39m), [33mAtom[39m([32m"Hello"[39m)))
)
[36mv5[39m: [32mAndOperator[39m = [33mAndOperator[39m(
  [33mAtom[39m([32m"Hello"[39m),
  [33mKleeneStar[39m(
    [33mOrOperator[39m([33mAtom[39m([32m"Hello"[39m), [33mConcatenation[39m([33mAtom[39m([32m"Hello"[39m), [33mAtom[39m([32m"Hello"[39m)))
  )
)
```

### C (8 points)
Write down the representation of the regular expression in Scala. Your cell must define a term that should be called `finalAnswerC`.

$$ ( \text{"hello"} |  (\text{"scala"}; \text{"best"})^* )^*$$


<!-- grade_id: answer_3c, points: 0, grade: false, solution: true -->
```scala
val finalAnswerC = {
    // YOUR CODE HERE
  KleeneStar ( OrOperator(Atom("hello"), KleeneStar(Concatenation(Atom("scala"),Atom("best")))))
}
```

Output:
```
[36mfinalAnswerC[39m: [32mKleeneStar[39m = [33mKleeneStar[39m(
  [33mOrOperator[39m(
    [33mAtom[39m([32m"hello"[39m),
    [33mKleeneStar[39m([33mConcatenation[39m([33mAtom[39m([32m"scala"[39m), [33mAtom[39m([32m"best"[39m)))
  )
)
```

<!-- grade_id: tests_3c, points: 8, grade: true, solution: false -->
```scala
//BEGIN TEST
def munge(r: Regex): String = r match {
    case Atom(st) => "$A$"+st
    case Concatenation(r1, r2) => munge(r1)+"$C$"+munge(r2)
    case KleeneStar(r) => "$K$"+munge(r)+"$S$"
    case OrOperator(r1, r2) => munge(r1)+"_O_"+munge(r2)
    case AndOperator(r1,r2) => munge(r1)+"/\\"+ munge(r2)
}

assert(munge(finalAnswerC) == "$K$$A$hello_O_$K$$A$scala$C$$A$best$S$$S$")
passed(8)
//END TEST
```

Output:
```
*** Tests Passed (8 points) ***
defined [32mfunction[39m [36mmunge[39m
```

```scala

```