Before you turn this problem in, make sure everything runs as expected. 
  1. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and 
  2. Then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name below. 

---

# CSCI 3155: Assignment 9

Topics: 
- Type Checking
- Basics of Objects
- Generics

Readings: Notes posted on canvas.

__Name__: WRITE YOUR NAME HERE

```scala
// TEST HELPER
def passed(points: Int) {
    require(points >=0)
    if (points == 1) print(s"\n*** Tests Passed (1 point) ***\n")
    else print(s"\n*** Tests Passed ($points points) ***\n")
}
```

Output:
```
defined [32mfunction[39m [36mpassed[39m
```

## Problem 1 (20 points)

Suppose we wish to define an abstract class called `Shape` that is a geometric shape. It needs to have methods such as  `getCenter`,  and `translate`.

  - `getCenter` : Gets the center of the shape. For a polygon the center is the "average" of the coordinates of its vertices. For a circle/ellipse, the center is part of the description of the object (see below).
  - `translate`: Translates the shape along some given values of `xShift` and `yShift`.


We will define a trait called `WithCorners` that define shapes with corners. This trait will implement a function called `getVertices`.

There are many classes that inherit from `Shape` including `Polygon` and `Ellipse`. The classes `Triangle` and `Rectangle` inherit from `Polygon` and `Circle` inherits from `Ellipse`. 

Also the class `Polygon` and all its derived classes must mixin the trait `WithCorners` but the class `Ellipse` and its derived classes should not.

- class `Triangle(x0: (Double, Double), x1: (Double, Double), x2: (Double, Double)))` must have three class parameters representing the three corners of the triangle.
  - It must inherit from `Polygon`.
  - It must mixin the trait `WithCorners`.
  - It must override the `translate` method so that `translate` called on a `Triangle` returns an instance of a `Triangle`.
  
- class `Rectangle(lowerLeft: (Double, Double), height: Double, width: Double)` must have the class parameters representing the lower left coordinate, height and width.
  - It must inherit from `Polygon`.
  - It must mixin the trait `WithCorners`.
  - It must override the `translate` method so that `translate` called on a `Rectangle` returns an instance of a `Rectangle`.
  
- class `Circle(center: (Double, Double), rad: Double)` must have its center coordinate and radius as class paramters.
  - It must inherit from Ellipse.
  - It must override the `translate` methodd so that `translate` called on a `Circle` returns an instance of a `Circle`.





### Part A
Complete the definitions below so that your code compiles.

**Restrictions** No loops, mutables (var) or recursion. Please use functors `map`, `foldLeft` etc.. when possible.


<!-- grade_id: cell-f3c8a032a2efce9b, points: 0, grade: false, solution: true -->
```scala
abstract class Shape {
    def getCenter:(Double, Double)
    def translate(xShift: Double, yShift: Double): Shape
}

trait WithCorners {
    def getVertices: List[(Double, Double)]
}

class Ellipse(val center: (Double, Double), val axisLengths: (Double, Double)) extends Shape {
    //TODO: Finish the methods that need to be implemented.
    // YOUR CODE HERE
      override def getCenter:(Double, Double)=center
   override def translate(xShift: Double, yShift: Double): Ellipse=
    {
        new Ellipse((center._1+xShift,center._2+yShift),axisLengths)
    }
}

class Polygon(val listOfVerts: List[(Double, Double)]) extends Shape with WithCorners {
     assert(listOfVerts.length >= 1)
    // TODO: Finish the methods that need to be implemented.
    // YOUR CODE HERE
    override def getCenter:(Double, Double)=
    {
        val totalLength=listOfVerts.length
        val x=listOfVerts.foldLeft(0.0)((acc:Double,element:(Double,Double)) => acc+element._1)
        val y=listOfVerts.foldLeft(0.0)((acc:Double,element:(Double,Double)) => acc+element._2)
        return (x/totalLength,y/totalLength)
    }
   override def translate(xShift: Double, yShift: Double): Polygon=
    {
        new Polygon(listOfVerts.map(x=>(x._1+xShift,x._2+yShift)))
    }
    override def getVertices: List[(Double, Double)]=listOfVerts
}

//TODO: Complete definitions of triangle, rectangle and circle classes.
// YOUR CODE HERE
class Triangle(x0: (Double, Double), x1: (Double, Double), x2: (Double, Double))extends Polygon(List(x0,x1,x2)) with WithCorners
{
   override def translate(xShift: Double, yShift: Double): Triangle=
    {
        new Triangle((x0._1+xShift,x0._2+yShift),(x1._1+xShift,x1._2+yShift),(x2._1+xShift,x2._2+yShift))
    }
}
class Rectangle(lowerLeft: (Double, Double), height: Double, width: Double)extends 
Polygon(List(lowerLeft,(lowerLeft._1,lowerLeft._2+width),(lowerLeft._1+height,lowerLeft._2),(lowerLeft._1+height,lowerLeft._2+width) )) with WithCorners
{
    override def translate(xShift: Double, yShift: Double): Rectangle=
    {
        new Rectangle((lowerLeft._1+xShift,lowerLeft._2+yShift),height,width)
    }
}
class Circle(center: (Double, Double), rad: Double)extends Ellipse(center,(rad,rad)) 
{
     override def translate(xShift: Double, yShift: Double): Circle=
    {
        new Circle( (center._1+xShift,center._2+yShift),rad)
    }
}
```

Output:
```
defined [32mclass[39m [36mShape[39m
defined [32mtrait[39m [36mWithCorners[39m
defined [32mclass[39m [36mEllipse[39m
defined [32mclass[39m [36mPolygon[39m
defined [32mclass[39m [36mTriangle[39m
defined [32mclass[39m [36mRectangle[39m
defined [32mclass[39m [36mCircle[39m
```

<!-- grade_id: cell-e38ff7897d7788b7, points: 10, grade: true, solution: false -->
```scala
val el1 = new Ellipse((5.0, -3.5), (2.1, 1.3))
val el2 = el1.translate(3,4)
assert (el2.center == (8.0, 0.5 ), "Test failed: After translation, the center must be (8,0.5)")

val pl1 = new Polygon(List((0,0), (2,1)))
assert (pl1.getCenter == (1.0, 0.5), "center must be (1,0.5)")

passed(10)
```

Output:
```
*** Tests Passed (10 points) ***
[36mel1[39m: [32mEllipse[39m = ammonite.$sess.cmd54$Helper$Ellipse@6d295cdb
[36mel2[39m: [32mEllipse[39m = ammonite.$sess.cmd54$Helper$Ellipse@45a5c8c3
[36mpl1[39m: [32mPolygon[39m = ammonite.$sess.cmd54$Helper$Polygon@4301179c
```

<!-- grade_id: cell-86f964f4436e9d8b, points: 5, grade: true, solution: false -->
```scala
val tri = new Triangle( (0,0), (2,0), (1,1))
val tri2 = tri.translate(1,1)
assert (tri2.isInstanceOf[Triangle])
print(tri2.getCenter)
assert(tri2.getCenter._1 == 2.0)
assert(1.3333333 <= tri2.getCenter._2 && tri2.getCenter._2 <= 1.333333334)
passed(5)
```

Output:
```
(2.0,1.3333333333333333)
*** Tests Passed (5 points) ***
[36mtri[39m: [32mTriangle[39m = ammonite.$sess.cmd54$Helper$Triangle@5f328922
[36mtri2[39m: [32mTriangle[39m = ammonite.$sess.cmd54$Helper$Triangle@3b2fcb90
```

<!-- grade_id: cell-0278b6bfaa79e728, points: 5, grade: true, solution: false -->
```scala
val rect = new Rectangle((-1,-1), 2, 2)
val listOfVerts = rect.getVertices
assert(listOfVerts.length == 4)
assert (rect.getCenter == (0,0))
val rect2: Rectangle = rect.translate(1,1)
assert(rect2.getCenter == (1,1))
passed(5)
```

Output:
```
*** Tests Passed (5 points) ***
[36mrect[39m: [32mRectangle[39m = ammonite.$sess.cmd54$Helper$Rectangle@4ca780b
[36mlistOfVerts[39m: [32mList[39m[([32mDouble[39m, [32mDouble[39m)] = [33mList[39m(
  ([32m-1.0[39m, [32m-1.0[39m),
  ([32m-1.0[39m, [32m1.0[39m),
  ([32m1.0[39m, [32m-1.0[39m),
  ([32m1.0[39m, [32m1.0[39m)
)
[36mrect2[39m: [32mRectangle[39m = ammonite.$sess.cmd54$Helper$Rectangle@33e00bbd
```

<!-- grade_id: cell-4e0c0f9982403849, points: 5, grade: true, solution: false -->
```scala
val circ = new Circle((1,1), 3)
val new_circ: Circle = circ.translate(2,2)
assert (new_circ.center == (3,3))
passed(5)
```

Output:
```
*** Tests Passed (5 points) ***
[36mcirc[39m: [32mCircle[39m = ammonite.$sess.cmd54$Helper$Circle@359abd68
[36mnew_circ[39m: [32mCircle[39m = ammonite.$sess.cmd54$Helper$Circle@290f6443
```

## Problem 2 (20 points): Objects in Scala

### A (12 points)
Consider the class below for an employee record in an organization that has three fields `name`, `eid` (employee id) and `salary`.

1) As given, scala will complain that class employee cannot extend the abstract class People. Add the missing methods to enable the class Employee to properly extend People

2) Write a factory pattern that takes in a formatted string of the form "John Doe,29585,33110" consisting of a string and two numbers, to extract the name and salaries for each employee. You may want to use `scala.util.matching.Regex` [https://www.scala-lang.org/api/2.9.2/scala/util/matching/Regex.html]
If the input fails to match the pattern, an `IllegalArgumentException` must be thrown

Companion objects are covered in scala book Chapter 4.3. Also lookup here: https://alvinalexander.com/scala/factory-pattern-in-scala-design-patterns

<!-- grade_id: cell-f80a491b407a7ab7, points: 0, grade: false, solution: true -->
```scala
import scala.util.matching.Regex

abstract class People {
    /* I need the ability to convert every person to a string */
    def toString: String
    /* I need the ability to check if two people are the same */
    def equals (p: People): Boolean
}

/* 
  Note that for an employee to be equal to an object of type People, that object
  must be an employee with matching eid number 
*/

class Employee( val name: String, val eid: Int, var salary: Int) extends People {
    // YOUR CODE HERE
   override def toString=name
    override  def equals(people:People):Boolean=
    {
        if(people.isInstanceOf[Employee] && people.asInstanceOf[Employee].eid==eid)
        {
         true
        }
        else
        {
             false
        }
    }
}

object Employee {
    def apply(formattedInput: String): Employee = {
        /* 
        Input format must be 
        (full name with possible spaces)[optional whitespaces],[optional whitespaces](employee ID number)[optional whitespaces],[optional whitespaces](employee salary)
        */
        // YOUR CODE HERE
        val form="""([a-zA-Z\s]+?)\s*,\s*(\d+)\s*,\s*(\d+)""".r //deal with formatted,like extra space
        formattedInput match
        {
            case form(name,eid,salary)=>new Employee(name,eid.toInt,salary.toInt)
            case _ => throw new IllegalArgumentException
        }
        
    }
}
```

Output:
```
[32mimport [39m[36mscala.util.matching.Regex

[39m
defined [32mclass[39m [36mPeople[39m
defined [32mclass[39m [36mEmployee[39m
defined [32mobject[39m [36mEmployee[39m
```

<!-- grade_id: cell-e1ec3bf75d6f6de5, points: 4, grade: true, solution: false -->
```scala
//BEGIN TEST
val e1 = Employee("Jonathan Doe, 12893, 39813")
println(e1.toString)
assert(e1.name == "Jonathan Doe", "Failed")
assert(e1.eid == 12893, "Failed")
assert(e1.salary == 39813, "Failed")
assert(e1.equals(new Employee("whoknows", 12893, 39813)))
passed(4)
//END TEST
```

Output:
```
Jonathan Doe

*** Tests Passed (4 points) ***
[36me1[39m: [32mEmployee[39m = Jonathan Doe
```

<!-- grade_id: cell-9d823f9842fd95a7, points: 4, grade: true, solution: false -->
```scala
//BEGIN TEST
try {
    val e2 = Employee("Badly Formatted String, 29812, $3102")
    assert(false, "Failed, the string is badly formatted but your code is OK with it.")
} catch {
    case e => println(s"Expected behavior seen")
}
passed(4)
//END TEST
```

Output:
```
Expected behavior seen

*** Tests Passed (4 points) ***
```

<!-- grade_id: cell-00d1907d97c7da69, points: 4, grade: true, solution: false -->
```scala
//BEGIN TEST
val e1 = Employee("Janet Maria Catherine Doe    , 12894   , 121331")
println(e1.toString)
assert(e1.name == "Janet Maria Catherine Doe", "Failed")
assert(e1.eid == 12894, "Failed")
assert(e1.salary == 121331, "Failed")
assert(e1.equals(new Employee("random", 12894, 121332)))
passed(4)
//END TEST
```

Output:
```
Janet Maria Catherine Doe

*** Tests Passed (4 points) ***
[36me1[39m: [32mEmployee[39m = Janet Maria Catherine Doe
```

### B (8 points)
We wish to add a new class of people to our system. These are "Contractor" with the information associated should involve their name (string), contract ID (Integer) and contract supervisor (Employee).

Write a new class named `Contractor` whose constructor should take in the name of the person (String), their ID number (Integer) and Supervisor (Employee) and extend from People. 

A contractor object is equal to an object p of type People, only if p is actually an instance of contractor, and if their names are the same and they have the same contract ID number.

<!-- grade_id: cell-97ab958367cfbd0a, points: 0, grade: false, solution: true -->
```scala
// YOUR CODE HERE

class Contractor( val name: String, val cid: Int, var supervisor: Employee) extends People
{
    override def toString=name
    //use p as people, p has type People
    override def equals(p:People):Boolean=
    {
        if(p.isInstanceOf[Contractor] && p.asInstanceOf[Contractor].name==name&&p.asInstanceOf[Contractor].cid==cid )
        {
             true
        }
        else
        {
             false
        }
    }
}
```

Output:
```
defined [32mclass[39m [36mContractor[39m
```

<!-- grade_id: cell-82cc00e175a12efe, points: 4, grade: true, solution: false -->
```scala
//BEGIN TEST
val e1 = Employee("Jonathan Doe, 12893, 39813")
val c1 = new Contractor("John Contract Worker", 12093, e1)
assert( !c1.equals(e1), "Failed")
assert( c1.equals(c1), "Failed")
passed(4)
//END TEST
```

Output:
```
*** Tests Passed (4 points) ***
[36me1[39m: [32mEmployee[39m = Jonathan Doe
[36mc1[39m: [32mContractor[39m = John Contract Worker
```

<!-- grade_id: cell-62fa2e0b66f58862, points: 4, grade: true, solution: false -->
```scala
//BEGIN TEST
val e1 = Employee("Jonathan Doe, 12893, 39813")
val e2 = Employee("Jane Doe, 12894, 43814")
val c1 = new Contractor("John Contract Worker", 12093, e2)
val c2 = new Contractor("John Contract Worker", 12093, e1)
assert( c1.equals(c2), "Failed")
passed(4)
//END TEST
```

Output:
```
*** Tests Passed (4 points) ***
[36me1[39m: [32mEmployee[39m = Jonathan Doe
[36me2[39m: [32mEmployee[39m = Jane Doe
[36mc1[39m: [32mContractor[39m = John Contract Worker
[36mc2[39m: [32mContractor[39m = John Contract Worker
```

## Problem 3 (10 points): Traits in Scala

For this problem, we have defined two traits: `NumberOfLegs` that helps us define how many legs a given animal has and `WarmBlooded` that applies to warm blooded animals.

```scala
abstract class Animal 

trait NumberOfLegs {
    val nLegs: Int
    def getNumberOfLegs: Int = nLegs
}

trait WarmBlooded extends Animal {
    val bodyTempMaintained: Double
    def getBodyTemp: Double = bodyTempMaintained 
}
```

Output:
```
defined [32mclass[39m [36mAnimal[39m
defined [32mtrait[39m [36mNumberOfLegs[39m
defined [32mtrait[39m [36mWarmBlooded[39m
```

### A (5 points)

Define a class `Human` that inherits from `Animal` and mixes in the traits `NumberOfLegs` with `nLegs = 2` and `WarmBlooded` with `bodyTempMaintained = 98`. The class `Human` must take in a parameter called `name` of type `String` and implement a `getName` method without any parameters.

<!-- grade_id: cell-e7a4a9275bd24e24, points: 0, grade: false, solution: true -->
```scala
// YOUR CODE HERE
class Human(val name:String,val nLegs:Int=2,val bodyTempMaintained:Double=98 )extends 
WarmBlooded with NumberOfLegs{
    def getName=name
}
```

Output:
```
defined [32mclass[39m [36mHuman[39m
```

<!-- grade_id: cell-f9b7ae8902326372, points: 5, grade: true, solution: false -->
```scala
//BEGIN TEST
val t1 = new Human("Jane Smith")
assert(t1.getNumberOfLegs == 2, "Your human does not have two legs")
assert(t1.bodyTempMaintained == 98.0, "Your human does not maintain a body temp of 98")
assert(t1.getBodyTemp == 98.0, "Your human's getBodyTemp Function is not working")
assert(t1.getName == "Jane Smith", "Your human's name is not setting correctly")
passed(5)
//END TEST
```

Output:
```
*** Tests Passed (5 points) ***
[36mt1[39m: [32mHuman[39m = ammonite.$sess.cmd67$Helper$Human@49587c1e
```

### B (5 points)
Define a class named `Table` that mixes in the trait `NumberOfLegs` with `nLegs = 4`. 

_Note:_ `Table` cannot mixin the trait `WarmBlooded` (make sure you understand why that is as a quick check of concepts).

<!-- grade_id: cell-8145df690b6f6588, points: 0, grade: false, solution: true -->
```scala
// YOUR CODE HERE
class Table(val nLegs:Int =4) extends NumberOfLegs
```

Output:
```
defined [32mclass[39m [36mTable[39m
```

<!-- grade_id: cell-d34b29b062b93894, points: 5, grade: true, solution: false -->
```scala
//BEGIN TEST
val tbl = new Table()
assert(tbl.getNumberOfLegs == 4, "A Table must have four legs")
passed(5)
//END TEST
```

Output:
```
*** Tests Passed (5 points) ***
[36mtbl[39m: [32mTable[39m = ammonite.$sess.cmd5$Helper$Table@33fb2f13
```

```scala

```