Before you turn this problem in, make sure everything runs as expected. 
  1. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and 
  2. Then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name below. 

---

# CSCI 3155: Assignment 4

Topics Covered: Map, foldLeft and filter functions to replace loops.

__Name__: WRITE YOUR NAME HERE

```scala
// TEST HELPER: PLEASE run this cell first
def passed(points: Int) {
    require(points >=0)
    if (points == 1) print(s"Tests Passed (1 point)")
    else print(s"Tests Passed ($points points)") 
}
```

Output:
```
defined [32mfunction[39m [36mpassed[39m
```

## Problem 1: map, filter, reduce on containers (40 points)

Solve the problems using a combination of map, filter and foldLeft/foldRight opertions over lists. 
 - The use of mutables, recursion, For/While loops is forbidden for all problems.
 - Each problem has very specific set of allowed LIST API functions. Those are the only functions you can use. Other API functions are not allowed. 
 - There are no restrictions on use of integer operations (addition, subtraction, division, mod etc..) or use of if-then-else, case-pattern matching and other scala constructs (however once again no loops of any kind and no use of var).
 - There are no restrictions on the use of named "helper" function vs anonymous functions.
 - Avoid the use of return in your anonymous function -- it may cause errors that are hard to debug.


### A: Compute a list of element-wise maximums of two lists of numbers (5 points).
Write a function `elementWiseMax` that takes two lists `lstA` and `lstB` of integers. The lists are given to be the same size. You need to compare the lists element-wise, i.e., compare $i$th element of `lstA` to $i$th element of `lstB`, and return a new list containing the maximum of compared elements at each index. 

$$\texttt{elementWiseMax}([a_0,a_1, \ldots, a_n], ~[b_0,b_1, \ldots, b_n]) = [\texttt{max}(a_0,b_0),\texttt{max}(a_1,b_1), \ldots, \texttt{max}(a_n,b_n)]$$.

Maximum of two integers is defined as following: 
$$\texttt{max}(a_i,b_i) ~=~ \texttt{if}~ b_i > a_i ~\texttt{then}~ b_i ~\texttt{else}~ a_i$$. 

**Example**

`elementWiseMax(List(1,3,5,4), List(-1,5,-2,7))` should return `List(1,5,5,7)`.

**Restrictions**

 - List API functions that you are allowed to use: `zip` and `map`. 
 - Use of any other List API function is not allowed.
 - You are __not__ allowed to use __var__, any kind of loops or recursion.

You can look the list API documentation for scala to find out more about these functions.
Post/Search on piazza if you are unsure. 



<!-- grade_id: cell-d07d79524c90fd4e, points: 0, grade: false, solution: true -->
```scala
def elementWiseMax(lstA: List[Int], lstB: List[Int]): List[Int] =
    
    {
        require(lstA.length == lstB.length)
    // YOUR CODE HERE
        def helper(a:Int,b:Int):Int=if(a>b)a else b
    val lst2:List[Int]= (lstA zip lstB).map(x=> helper(x._1,x._2) )
        lst2
        
    }

```

Output:
```
defined [32mfunction[39m [36melementWiseMax[39m
```

<!-- grade_id: cell-054acf67575fc99c, points: 5, grade: true, solution: false -->
```scala
// BEGIN TEST

val t1 = elementWiseMax(List(1,2,3), List(4,5,6))
assert(t1 == List(4,5,6), "Test 1 failed")

val t2 = elementWiseMax(List(1,4,0), List(2,1,5))
assert(t2 == List(2,4,5), "Test 2 failed")

val t3 = elementWiseMax(List(), List())
assert(t3 == List(), "Test 3 failed")

val t4 = elementWiseMax(List(1, 0, 3, -1, -2), List(0, 4, 1, 0, -3))
assert(t4 == List(1, 4, 3, 0, -2), "Test 4 failed")

passed(5)
// END TEST
```

Output:
```
Tests Passed (5 points)[36mt1[39m: [32mList[39m[[32mInt[39m] = [33mList[39m([32m4[39m, [32m5[39m, [32m6[39m)
[36mt2[39m: [32mList[39m[[32mInt[39m] = [33mList[39m([32m2[39m, [32m4[39m, [32m5[39m)
[36mt3[39m: [32mList[39m[[32mInt[39m] = [33mList[39m()
[36mt4[39m: [32mList[39m[[32mInt[39m] = [33mList[39m([32m1[39m, [32m4[39m, [32m3[39m, [32m0[39m, [32m-2[39m)
```

### B: LCM of all Denominators (7 points).

You are given a list of pairs of integers that are supposed to represent fractions. You can assume that all fractions are positive and already in their lowest terms. This problem asks you to compute the LCM of the denominators.

Eg., Input: `List((1,2), (3,4), (8,9), (15,24))`
The denominators are `2, 4, 9, 24` and their LCM is `72`, which is the answer your function should return.

For the empty list as input, your code must return the answer 1.

You should assume that all the denominators are positive.

**Restrictions**

  - You are allowed to use the provided `lcm` function that computes LCM of two numbers.
  - Allowed list API functions: `map`, `filter`, `foldLeft`, and `foldRight`. 
  - No other list API functions allowed.
  - No __var__, no loops and no recursive calls other than the call made for LCM.


<!-- grade_id: cell-332c55ac9d70c435, points: 0, grade: false, solution: true -->
```scala
import scala.annotation._

@tailrec
private def gcd(a: Int, b: Int):Int=if (b==0) a.abs else gcd(b, a%b)

def lcm(a: Int, b: Int)=(a*b).abs/gcd(a,b)

def lcmOfDenominators(lst: List[(Int, Int)]): Int = {
    // YOUR CODE HERE
    lst.map(x=> x._2).foldLeft(1)( (acc,x)=> lcm(acc,x))
   
}
```

Output:
```
[32mimport [39m[36mscala.annotation._

[39m
defined [32mfunction[39m [36mlcm[39m
defined [32mfunction[39m [36mlcmOfDenominators[39m
```

<!-- grade_id: cell-b37de6cd26982543, points: 7, grade: true, solution: false -->
```scala
// BEGIN TEST
val i1 = lcmOfDenominators(List((1,2), (3,4), (5,6),(8,9), (11,12)))
assert(i1 == 36, s"Test 1 failed -- expected answer is 36, you got $i1")
val i2 = lcmOfDenominators(List())
assert(i2 == 1, s"Test 2 failed -- empty list must trivially have lcm of 1")
val i3 = lcmOfDenominators(List((4,5),(5,9),(18,15)))
assert(i3 == 45, s"Test 2 failed -- your answer is $i3")
passed(7)
// END TEST
```

Output:
```
Tests Passed (7 points)[36mi1[39m: [32mInt[39m = [32m36[39m
[36mi2[39m: [32mInt[39m = [32m1[39m
[36mi3[39m: [32mInt[39m = [32m45[39m
```

## C: All List Elements That are Greater Than Their Index (6 points)

Write a list function that inputs a list of integers `lst` and outputs a list consisting of all those elements in `lst` whose values are __strictly greater__ than their indices. Indices of elements start at 0 and go upto `length of list` -1.
The order of the elements in the original list must be preserved in the returned list.

**Restrictions**

- You may use  these list API functions: `zip`, `zipWithIndex` (read documentation), `map`, `foldLeft`, `foldRight` and `filter`.
- No other list API functions are allowed.
- Use of loops, var, recursion not allowed.

#### Example

  - Input: List(1, 2, 2, -2, 4, 6,8)
  - Output: List(1, 2, 6, 8)
Note that elements 1 (index 0), 2 (index 1), 6 (index 5) and 8 (index 6) are returned. 



<!-- grade_id: cell-ec76364060e01d7b, points: 0, grade: false, solution: true -->
```scala
def listElementsGreaterThanIndex(lst: List[Int]): List[Int] = {
    // YOUR CODE HERE
   val lst2:List[Int]=lst.zipWithIndex.filter(x=> x._1>x._2 ).map(x=>x._1)
    lst2
}
```

Output:
```
defined [32mfunction[39m [36mlistElementsGreaterThanIndex[39m
```

<!-- grade_id: cell-bbaaea1a391dc9de, points: 6, grade: true, solution: false -->
```scala
val lst1 = List(1,2,2,-2,4,6,8)
val res1 = listElementsGreaterThanIndex(lst1)
assert( res1 == List(1,2,6,8))


val lst2 = List(-1,-2,-3,-4,-5,-6)
val res2 = listElementsGreaterThanIndex(lst2)
assert( res2.length == 0)

val lst3 = List(1,2,3,4,5,6,7,8)
val res3 = listElementsGreaterThanIndex(lst3)
assert( res3 == lst3)

val lst4 = List(10, 9, 8, 7, 6, 5, 4, 3, 2, 1, -1, -2, -3)
val res4 = listElementsGreaterThanIndex(lst4)
assert (res4 == List(10, 9, 8, 7, 6))

passed(6)

```

Output:
```
Tests Passed (6 points)[36mlst1[39m: [32mList[39m[[32mInt[39m] = [33mList[39m([32m1[39m, [32m2[39m, [32m2[39m, [32m-2[39m, [32m4[39m, [32m6[39m, [32m8[39m)
[36mres1[39m: [32mList[39m[[32mInt[39m] = [33mList[39m([32m1[39m, [32m2[39m, [32m6[39m, [32m8[39m)
[36mlst2[39m: [32mList[39m[[32mInt[39m] = [33mList[39m([32m-1[39m, [32m-2[39m, [32m-3[39m, [32m-4[39m, [32m-5[39m, [32m-6[39m)
[36mres2[39m: [32mList[39m[[32mInt[39m] = [33mList[39m()
[36mlst3[39m: [32mList[39m[[32mInt[39m] = [33mList[39m([32m1[39m, [32m2[39m, [32m3[39m, [32m4[39m, [32m5[39m, [32m6[39m, [32m7[39m, [32m8[39m)
[36mres3[39m: [32mList[39m[[32mInt[39m] = [33mList[39m([32m1[39m, [32m2[39m, [32m3[39m, [32m4[39m, [32m5[39m, [32m6[39m, [32m7[39m, [32m8[39m)
[36mlst4[39m: [32mList[39m[[32mInt[39m] = [33mList[39m([32m10[39m, [32m9[39m, [32m8[39m, [32m7[39m, [32m6[39m, [32m5[39m, [32m4[39m, [32m3
```

## D ( 6 points) : All List Indices whose Elements are Divisible by 3.

Write a function that inputs a list of integers and returns all the indices in the original list whose elements are divisible by 3. Note that indices start from 0. 

Your code must return the indices sorted in ascending order.

**Restrictions**
- Allowed List API functions: `zipWithIndex`, `zip`, `filter`, `map`, `foldLeft`, and `foldRight`.
- Not allowed: all other list API functions.
- Loops of any form, recursion and use of var not allowed.

### Example

  - Input: List(1, **3**, 5, 7, **6, 9, 12**, 11)
  - Output: List(1, 4, 5, 6)
  
  Note that the returned list is a list of integers whose entries correspond to the indices in the original list.



<!-- grade_id: cell-2399c480ddc85d87, points: 0, grade: false, solution: true -->
```scala
def allIndicesDivisibleBy3(lst: List[Int]): List[Int] = {
    // YOUR CODE HERE
   val lst2:List[Int]=lst.zipWithIndex.filter(x=>x._1%3==0).map(x=>x._2)
    lst2
}
```

Output:
```
defined [32mfunction[39m [36mallIndicesDivisibleBy3[39m
```

<!-- grade_id: cell-ebebb499fc5f05d4, points: 6, grade: true, solution: false -->
```scala
//BEGIN TESTS
val lst1 = List(1,3,5,7,6,9,12,11)
val res1 = allIndicesDivisibleBy3(lst1)
assert(res1 == List(1,4,5,6))

val lst2 = List(1,5,7,9,11,13,15)
val res2 = allIndicesDivisibleBy3(lst2)
assert(res2 == List(3,6))


val lst3 = List(1,5,7,11,13)
val res3 = allIndicesDivisibleBy3(lst3)
assert(res3.length == 0)


val lst4 = List(-3,-6,-9,-12,-15,-18,-21,-24)
val res4 = allIndicesDivisibleBy3(lst4)
assert(res4 == List(0,1,2,3,4,5,6,7))


//END TESTS
```

Output:
```
[36mlst1[39m: [32mList[39m[[32mInt[39m] = [33mList[39m([32m1[39m, [32m3[39m, [32m5[39m, [32m7[39m, [32m6[39m, [32m9[39m, [32m12[39m, [32m11[39m)
[36mres1[39m: [32mList[39m[[32mInt[39m] = [33mList[39m([32m1[39m, [32m4[39m, [32m5[39m, [32m6[39m)
[36mlst2[39m: [32mList[39m[[32mInt[39m] = [33mList[39m([32m1[39m, [32m5[39m, [32m7[39m, [32m9[39m, [32m11[39m, [32m13[39m, [32m15[39m)
[36mres2[39m: [32mList[39m[[32mInt[39m] = [33mList[39m([32m3[39m, [32m6[39m)
[36mlst3[39m: [32mList[39m[[32mInt[39m] = [33mList[39m([32m1[39m, [32m5[39m, [32m7[39m, [32m11[39m, [32m13[39m)
[36mres3[39m: [32mList[39m[[32mInt[39m] = [33mList[39m()
[36mlst4[39m: [32mList[39m[[32mInt[39m] = [33mList[39m([32m-3[39m, [32m-6[39m, [32m-9[39m, [32m-12[39m, [32m-15[39m, [32m-18[39m, [32m-21[39m, [32m-24[39m)
[36mres4[39m: [32mList[39m[[32mInt[39m] = [33mList[39m([32m0[39m, [32m1[39m, [32m2
```

## Problem 2E (7 points) : Convert a List of Strings into List of Tuples

Write a function that achieves the following.

Input: List of String

Output: List of 4-tuples: wherein each tuple consists of the corresponding string in the input list, its length, its reverse and its index.

See example below. Order of elements in the original list must be preserved.


### Example

  - Input: List("Que", "Sera", "Sera", "Whatever", "Will", "Be")
  - Output: List( ( "Que", 3, "euQ", 0), ("Sera", 4, "areS", 1), ("Sera", 4, "areS", 2), ("Whatever", 8, "revetahW", 3), ("Will", 4, "lliW", 4), ("Be", 2, "eB", 5)) )


**Restrictions**
- You are allowed to use `map`, `filter`, `foldLeft`, `foldRight`, `zip` and `zipWithIndex`.
- All other list API functions are disallowed.
- String API functions: `length` and `reverse` are allowed.
- All other string API functions disallowed.
- Use of var, recursion or any kind of loop is forbidden.



<!-- grade_id: cell-936623db4d48dc12, points: 0, grade: false, solution: true -->
```scala
def convertListToListTuples(lst: List[String]): List[(String, Int, String, Int)] = {
    // YOUR CODE HERE
      val lst2=lst.foldLeft [List[(String, Int, String, Int)]] (List()) (( acc,x) => acc++List((x,x.length,x.reverse,acc.size) ))
    lst2
}
```

Output:
```
defined [32mfunction[39m [36mconvertListToListTuples[39m
```

<!-- grade_id: cell-c920ce6277c82861, points: 7, grade: true, solution: false -->
```scala
//BEGIN TESTS
def loopTestListOfTuples (lst: List[String]): List[(String, Int, String, Int)] = {
    var count = 0
    var ret_list:List[(String, Int, String, Int)]  = List()
    for (elt <- lst){
        ret_list = ret_list ++ List((elt, elt.length, elt.reverse, count))
        count = count + 1
    }
    ret_list
}

val lst1 = List("Que", "Sera", "Sera", "Whatever", "Will", "Be")
val expected1 = loopTestListOfTuples(lst1)
assert(convertListToListTuples(lst1) == expected1)

val lst2 = List("All", "The", "World", "Is", "A", "Stage")
val expected2 = loopTestListOfTuples(lst2)
assert(convertListToListTuples(lst2) == expected2)


val lst3 = List("Cry", "Havoc!", "And", "Let", "Loose", "The", "Dogs", "Of", "War")
val expected3 = loopTestListOfTuples(lst3)
assert(convertListToListTuples(lst3) == expected3)


val lst4 = List("To", "Be", "Or", "Not", "to", "be", "is", "the", "question")
val expected4 = loopTestListOfTuples(lst4)
assert(convertListToListTuples(lst4) == expected4)

// END TESTS
passed(7)

```

Output:
```
Tests Passed (7 points)defined [32mfunction[39m [36mloopTestListOfTuples[39m
[36mlst1[39m: [32mList[39m[[32mString[39m] = [33mList[39m([32m"Que"[39m, [32m"Sera"[39m, [32m"Sera"[39m, [32m"Whatever"[39m, [32m"Will"[39m, [32m"Be"[39m)
[36mexpected1[39m: [32mList[39m[([32mString[39m, [32mInt[39m, [32mString[39m, [32mInt[39m)] = [33mList[39m(
  ([32m"Que"[39m, [32m3[39m, [32m"euQ"[39m, [32m0[39m),
  ([32m"Sera"[39m, [32m4[39m, [32m"areS"[39m, [32m1[39m),
  ([32m"Sera"[39m, [32m4[39m, [32m"areS"[39m, [32m2[39m),
  ([32m"Whatever"[39m, [32m8[39m, [32m"revetahW"[39m, [32m3[39m),
  ([32m"Will"[39m, [32m4[39m, [32m"lliW"[39m, [32m4[39m),
  ([32m"Be"[39m, [32m2[39m, [32m"eB"[39m, [32m5[39m)
)
[36mlst2[39m: [32mList[39m[[32mString[39m] = [33mList[39m([32m"All"[39m, [32m"The"[39m, [32m"World"[39m, [32m"Is"[39m, [32m"A"[39m, [32m"Stage"[39m)
[36mexpected2[39m: [32mList[39m[([32mString[39m,
```

### F (9 points): Convert a list into an indexed list.

Write a function to convert a list of strings into an indexed list of strings.  Indices for this particular problem must start at 1.

**Example**

__Input__ List("hello", "world", "my", "cat", "is", "grumpy", "today")

__Output__  List( ("hello", 1), ("world", 2), ("my", 3), ("cat", 4), ("is", 5), ("grumpy", 6), ("today", 7) )


**Restrictions**

- You are allowed to use just the basic list operations such as cons of an element to a list (::), concatenation of two lists (++, or ::: operators) and list length operations.
- List API functions `length`, `reverse`, `map`, `filter`, `foldLeft` and `foldRight` are allowed.
- No other list API functions (specifically use of the function `zipWithIndex` or `zip` will result in 0 points for the problem).
- __var__, __loops__ or __recursion__ are not allowed.

**Hint**

- Use `foldLeft` that accumulate a list of tuples, wherein each tuple has a string and a number.
- The size of the accumulator + 1 provides you with the "index" of the current element.


<!-- grade_id: cell-0c984f492b25b8b4, points: 0, grade: false, solution: true -->
```scala

def makeIndexedList(lst:List[String]): List[(String, Int)] = {
    // YOUR CODE HERE
   val lst2=lst.foldLeft [List[(String,Int)]] (List()) (( acc,x) => acc++List((x,1+acc.size)) )
    lst2
}
```

Output:
```
defined [32mfunction[39m [36mmakeIndexedList[39m
```

<!-- grade_id: cell-8c49d31044f1e0ae, points: 9, grade: true, solution: false -->
```scala
// BEGIN TEST
val t1 = makeIndexedList(List("hello"))
assert(t1 == List(("hello", 1)), s"Test 1 failed - your code returned $t1")

val t2 = makeIndexedList(List("hello", "world"))
assert(t2 == List(("hello", 1), ("world", 2)), s"Test 2 failed - your code returned $t2")

val t3 = makeIndexedList(Nil)
assert(t3 == Nil, s"Test 3 failed - your code returned $t3")

val t4 = makeIndexedList(List("a","b","c","d","e"))
assert(t4 == List(("a",1), ("b",2), ("c",3), ("d",4), ("e",5)), s"Test 4 failed - your code returned $t4")

passed(9)
// END TEST
```

Output:
```
Tests Passed (9 points)[36mt1[39m: [32mList[39m[([32mString[39m, [32mInt[39m)] = [33mList[39m(([32m"hello"[39m, [32m1[39m))
[36mt2[39m: [32mList[39m[([32mString[39m, [32mInt[39m)] = [33mList[39m(([32m"hello"[39m, [32m1[39m), ([32m"world"[39m, [32m2[39m))
[36mt3[39m: [32mList[39m[([32mString[39m, [32mInt[39m)] = [33mList[39m()
[36mt4[39m: [32mList[39m[([32mString[39m, [32mInt[39m)] = [33mList[39m(([32m"a"[39m, [32m1[39m), ([32m"b"[39m, [32m2[39m), ([32m"c"[39m, [32m3[39m), ([32m"d"[39m, [32m4[39m), ([32m"e"[39m, [32m5[39m))
```

### G: Partitioning a list (5 points)

Given a list of numbers, partition the list into a pair of lists where the first list contains only the even numbers and the second list contains only the odd numbers. Both the output lists must contain the elements in the same order as they occur in the input list. The result of partition an empty list is a pair of empty lists.

Example:

Input: List(1, 2, 8, 7, 5)

Output: (List(2,8), List(1,7,5))

**Restrictions**

- You are allowed to use just the basic list operatios such as cons of an element to a list (::) and concatenation of two lists (++, or ::: operators). 
- List API functions  `map`, `filter`, `foldLeft` and `foldRight` but no other list API functions. 
- Do not use __var__, __loops__ or __recursion__.


<!-- grade_id: cell-61ba46072b910129, points: 0, grade: false, solution: true -->
```scala
def oddEvenPartition(lst: List[Int]): (List[Int], List[Int]) = {
    // YOUR CODE HERE
   val lst2:(List[Int], List[Int])=(lst.filter(x=>x%2==0) ,lst.filter(x=>x%2!=0))
    lst2
   
}
```

Output:
```
defined [32mfunction[39m [36moddEvenPartition[39m
```

<!-- grade_id: cell-b766aa2568f6e251, points: 5, grade: true, solution: false -->
```scala
// BEGIN TEST
val x = oddEvenPartition(List(1, 2, 8, 7, 5))
assert(x._1 == List(2,8) && x._2 == List(1,7,5), s"Test 1 failed -- your code returned $x")

val y = oddEvenPartition(List(23, 18, 20, 15))
assert(y._1 == List(18,20) && y._2 == List(23,15), s"Test 2 failed -- your code returned $y")

val z = oddEvenPartition(List(14, 12, 18, 20))
assert(z._1 == List(14, 12, 18, 20) && z._2 == List(), s"Test 3 failed -- your code returned $z")

val w = oddEvenPartition(List())
assert(w._1 == List() && w._2 == List(), s"Test 4 failed -- your code returned $w")

passed(5)
// END TEST
```

Output:
```
Tests Passed (5 points)[36mx[39m: ([32mList[39m[[32mInt[39m], [32mList[39m[[32mInt[39m]) = ([33mList[39m([32m2[39m, [32m8[39m), [33mList[39m([32m1[39m, [32m7[39m, [32m5[39m))
[36my[39m: ([32mList[39m[[32mInt[39m], [32mList[39m[[32mInt[39m]) = ([33mList[39m([32m18[39m, [32m20[39m), [33mList[39m([32m23[39m, [32m15[39m))
[36mz[39m: ([32mList[39m[[32mInt[39m], [32mList[39m[[32mInt[39m]) = ([33mList[39m([32m14[39m, [32m12[39m, [32m18[39m, [32m20[39m), [33mList[39m())
[36mw[39m: ([32mList[39m[[32mInt[39m], [32mList[39m[[32mInt[39m]) = ([33mList[39m(), [33mList[39m())
```

```scala

```