Before you turn this problem in, make sure everything runs as expected. 
  1. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and 
  2. Then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name below. 

---

# CSCI 3155: Assignment 3 

Topics covered:Abstract Syntax of Programming Languages, and Operations on Inductive Definitions Using Pattern Matching.

Read quickly about the use of "Option" types for Problem 3A --  https://www.geeksforgeeks.org/scala-option/


```scala
// TEST HELPER
def passed(points: Int) {
    require(points >=0)
    if (points == 1) print(s"\n*** Tests Passed (1 point) ***\n")
    else print(s"\n*** Tests Passed ($points points) ***\n")
}
```

Output:
```
defined [32mfunction[39m [36mpassed[39m
```

## Problem 1 (15 points)

Develop a multi-switch statement for the `mython` (var) language developed in class. A multi-switch statement is a generalization of the popular switch statement to multiple expressions, i.e., instead of case-matching on one expression, we case-match on multiple expressions. The concrete syntax would look like this:

#### Example (Concrete Syntax)

~~~
switch (x, y, x-2*y, z) {
    case (1, 2.5, z, w) => { y = y + 10; z = y; }
    case (0, -1.5, 2.3, w-z+ 2*x) => { y = y - 10; z = 10}
    default => { z = 0 }
}
~~~

In this example, you are doing a multiway comparison on four expressions. The precise semantics of how such a statement would work are not relevant at this stage of the course. We ask you to help extend the _abstract syntax tree_ of Mython with these features.

Note that a valid multi-switch statement must enforce that if we are switching on k expressions at the same time, each case must provide k values to compare. We assume that this is taken care of by a check implemented by the parser. Ignore this consideration for the rest of the problem.


#### Grammar Specification for Multi-Switch

The syntax for the proposed multi-switch statement is given in terms of additional rules shown below that add to the existing definitions for the Mython language in the class notes on Inductive Definition (posted week 2).

$$\begin{array}{rcl}
\mathbf{Statement} & \rightarrow & \text{Switch}(\mathbf{Expr}*,  \mathbf{Case}* ) \\[5pt]
\mathbf{Case} & \rightarrow & \text{SwitchCase}(\mathbf{Expr}*, \mathbf{Statement}* ) \\[5pt]
& | & \text{DefaultCase}(\mathbf{Statement}*) \\
\end{array}$$

Note:
  - The  $*$ after a expression denotes "Kleene star" which is one or more instances of. 
  - As discussed in class and notes, we will use scala Lists to implement this.
  - The non-terminal __Case__ is being addded to the Mython grammar.
  - The non-terminal __Statement__ already exists and we are adding a new rule that makes _switch_ statements a new statement type.

The existing scala language definitions for `mython` are recalled for below from the class notes.


Augment the definitions by adding the extra definitions to support switch statements. Please use
the same constructor names as in the grammar: `Switch` for the switch statement and 
`SwitchCase` and `DefaultCase` for the two rules corresponding to the __Case__ specifications.

For your convenience, the definitions used in the class are given. Add your solution at the end.


<!-- grade_id: cell-a3a78387d2a93841, points: 0, grade: false, solution: true -->
```scala
/*-- Arithmetic Expressions --*/
sealed trait Expr
case class Const(d: Double) extends Expr 
case class Ident(s: String) extends Expr
case class Plus(e1: Expr, e2: Expr ) extends Expr
case class Minus(e1: Expr, e2: Expr) extends Expr
case class Mult(e1: Expr, e2: Expr) extends Expr
case class Div(e1: Expr, e2: Expr) extends Expr
case class Negate(e: Expr) extends Expr
case class Log(e: Expr) extends Expr
case class Exp(e: Expr) extends Expr
case class Sine(e: Expr) extends Expr
case class Cosine(e: Expr) extends Expr

/*-- Conditional Expressions --*/
sealed trait CondExpr
case object ConstTrue extends CondExpr
case object ConstFalse extends CondExpr
case class Geq(e1: Expr, e2: Expr) extends CondExpr
case class Leq(e1: Expr, e2: Expr) extends CondExpr
case class Eq(e1: Expr, e2: Expr) extends CondExpr
case class And(c1: CondExpr, c2: CondExpr) extends CondExpr
case class Or(c1: CondExpr, c2: CondExpr) extends CondExpr
case class Not(c: CondExpr) extends CondExpr

/*-- Declarations, Statements and such.. --*/
sealed trait Declaration
sealed trait Statement
sealed trait MythonProgram

case class Program(decls: List[Declaration], stmts: List[Statement], returnAtEnd: Expr) extends MythonProgram // We stripped the ReturnStmt tag since it is redundant
case class VarDecl(identifier: String, rhsExpr: Expr) extends Declaration
case class AssignStmt(identifier: String, rhsExpr: Expr) extends Statement
case class WhileStmt(cond: CondExpr, stmts: List[Statement]) extends Statement
case class IfThenElseStmt(cond: CondExpr, stmtsThen: List[Statement], stmtsElse: List[Statement]) extends Statement
case class ReturnStmt(retExpr: Expr) extends Statement

/*-- Write down the definitions of switch cases here --*/
// YOUR CODE HERE

sealed trait Case
case class SwitchCase(e:List[Expr],stmts:List[Statement]) extends Case
case class DefaultCase(stmts:List[Statement]) extends Case
case class Switch(e:List[Expr],c:List[Case]) extends Statement
```

Output:
```
defined [32mtrait[39m [36mExpr[39m
defined [32mclass[39m [36mConst[39m
defined [32mclass[39m [36mIdent[39m
defined [32mclass[39m [36mPlus[39m
defined [32mclass[39m [36mMinus[39m
defined [32mclass[39m [36mMult[39m
defined [32mclass[39m [36mDiv[39m
defined [32mclass[39m [36mNegate[39m
defined [32mclass[39m [36mLog[39m
defined [32mclass[39m [36mExp[39m
defined [32mclass[39m [36mSine[39m
defined [32mclass[39m [36mCosine[39m
defined [32mtrait[39m [36mCondExpr[39m
defined [32mobject[39m [36mConstTrue[39m
defined [32mobject[39m [36mConstFalse[39m
defined [32mclass[39m [36mGeq[39m
defined [32mclass[39m [36mLeq[39m
defined [32mclass[39m [36mEq[39m
defined [32mclass[39m [36mAnd[39m
defined [32mclass[39m [36mOr[39m
defined [32mclass[39m [36mNot[39m
defined [32mtrait[39m [36mDeclaration[39m
defined [32mtrait[39m [36mStatement[39m
defined [32mtrait[39m [36mMythonProgram[39m
defined [32mclass[39m [36mP
```

<!-- grade_id: cell-b4433298fac96b06, points: 5, grade: true, solution: false -->
```scala
//BEGIN TEST
val e1 = Plus(Ident("y"), Plus( Ident("z"), Exp( Minus(Ident("x"), Ident("y")) ) )) // y + z + exp(x - y)
val e2 = Div(Plus(Ident("x"), Ident("y")), Const(2)) // (x + y)/2 
val stmt1 = AssignStmt("x", e1) // x = y + z + exp(x - y)
// Let us make up the parts of the while statement
val cond1 = Leq(Ident("y"), Const(15.0f)) // y <= 15
// The assignment inside the while
val stmt2 = AssignStmt("y", Minus(Ident("y"), Ident("x"))) // y = y - x
// The condition for the if statement
val cond2 = Leq(Ident("x"), Const(0.0f)) // x <= 0.0
val stmt3 = AssignStmt("x", Minus(Const(0.0f), Ident("x"))) // x= 0 - x
val ifStmt = IfThenElseStmt(cond2, List(stmt3), List()) // if ( x <= 0) x = -x else nothing
val whileStmt = WhileStmt(cond1, List(stmt2, ifStmt)) // while loop 
val allstmts = List(stmt1, whileStmt)
val retExpr = Minus(Ident("y"), Ident("x")) // y - x to be returned

// If this goes through without an error, you passed the tests
val c1 = SwitchCase(List(Const(2.0), Const(3.0)), List(stmt1, stmt3))
val c3 = SwitchCase(List(Ident("y"), Ident("z")), List(stmt1, stmt2, whileStmt))

passed(5)
//END TEST
```

Output:
```
*** Tests Passed (5 points) ***
[36me1[39m: [32mPlus[39m = [33mPlus[39m(
  [33mIdent[39m([32m"y"[39m),
  [33mPlus[39m([33mIdent[39m([32m"z"[39m), [33mExp[39m([33mMinus[39m([33mIdent[39m([32m"x"[39m), [33mIdent[39m([32m"y"[39m))))
)
[36me2[39m: [32mDiv[39m = [33mDiv[39m([33mPlus[39m([33mIdent[39m([32m"x"[39m), [33mIdent[39m([32m"y"[39m)), [33mConst[39m([32m2.0[39m))
[36mstmt1[39m: [32mAssignStmt[39m = [33mAssignStmt[39m(
  [32m"x"[39m,
  [33mPlus[39m([33mIdent[39m([32m"y"[39m), [33mPlus[39m([33mIdent[39m([32m"z"[39m), [33mExp[39m([33mMinus[39m([33mIdent[39m([32m"x"[39m), [33mIdent[39m([32m"y"[39m)))))
)
[36mcond1[39m: [32mLeq[39m = [33mLeq[39m([33mIdent[39m([32m"y"[39m), [33mConst[39m([32m15.0[39m))
[36mstmt2[39m: [32mAssignStmt[39m = [33mAssignStmt[39m([32m"y"[39m, [33mMinus[39m([33mIdent[39m([32m"y"[39m), [33mIdent[39m([32m"x"[39m)))
[36mcond2[39m: [32mLeq[39m = [33mLeq
```

<!-- grade_id: cell-590403a05cbd399c, points: 5, grade: true, solution: false -->
```scala
//BEGIN TEST
val e1 = Plus(Ident("y"), Plus( Ident("z"), Exp( Minus(Ident("x"), Ident("y")) ) )) // y + z + exp(x - y)
val e2 = Div(Plus(Ident("x"), Ident("y")), Const(2)) // (x + y)/2 
val stmt1 = AssignStmt("x", e1) // x = y + z + exp(x - y)
// Let us make up the parts of the while statement
val cond1 = Leq(Ident("y"), Const(15.0f)) // y <= 15
// The assignment inside the while
val stmt2 = AssignStmt("y", Minus(Ident("y"), Ident("x"))) // y = y - x
// The condition for the if statement
val cond2 = Leq(Ident("x"), Const(0.0f)) // x <= 0.0
val stmt3 = AssignStmt("x", Minus(Const(0.0f), Ident("x"))) // x= 0 - x
val ifStmt = IfThenElseStmt(cond2, List(stmt3), List()) // if ( x <= 0) x = -x else nothing
val whileStmt = WhileStmt(cond1, List(stmt2, ifStmt)) // while loop 
val allstmts = List(stmt1, whileStmt)
val retExpr = Minus(Ident("y"), Ident("x")) // y - x to be returned
// Check if we can write a default case with an if statement
val c2 = DefaultCase(List(ifStmt))

passed(5)
//END TEST
```

Output:
```
*** Tests Passed (5 points) ***
[36me1[39m: [32mPlus[39m = [33mPlus[39m(
  [33mIdent[39m([32m"y"[39m),
  [33mPlus[39m([33mIdent[39m([32m"z"[39m), [33mExp[39m([33mMinus[39m([33mIdent[39m([32m"x"[39m), [33mIdent[39m([32m"y"[39m))))
)
[36me2[39m: [32mDiv[39m = [33mDiv[39m([33mPlus[39m([33mIdent[39m([32m"x"[39m), [33mIdent[39m([32m"y"[39m)), [33mConst[39m([32m2.0[39m))
[36mstmt1[39m: [32mAssignStmt[39m = [33mAssignStmt[39m(
  [32m"x"[39m,
  [33mPlus[39m([33mIdent[39m([32m"y"[39m), [33mPlus[39m([33mIdent[39m([32m"z"[39m), [33mExp[39m([33mMinus[39m([33mIdent[39m([32m"x"[39m), [33mIdent[39m([32m"y"[39m)))))
)
[36mcond1[39m: [32mLeq[39m = [33mLeq[39m([33mIdent[39m([32m"y"[39m), [33mConst[39m([32m15.0[39m))
[36mstmt2[39m: [32mAssignStmt[39m = [33mAssignStmt[39m([32m"y"[39m, [33mMinus[39m([33mIdent[39m([32m"y"[39m), [33mIdent[39m([32m"x"[39m)))
[36mcond2[39m: [32mLeq[39m = [33mLeq
```

<!-- grade_id: cell-73f9de24de44bb2f, points: 5, grade: true, solution: false -->
```scala
//BEGIN TEST
val e1 = Plus(Ident("y"), Plus( Ident("z"), Exp( Minus(Ident("x"), Ident("y")) ) )) // y + z + exp(x - y)
val e2 = Div(Plus(Ident("x"), Ident("y")), Const(2)) // (x + y)/2 
val stmt1 = AssignStmt("x", e1) // x = y + z + exp(x - y)
// Let us make up the parts of the while statement
val cond1 = Leq(Ident("y"), Const(15.0f)) // y <= 15
// The assignment inside the while
val stmt2 = AssignStmt("y", Minus(Ident("y"), Ident("x"))) // y = y - x
// The condition for the if statement
val cond2 = Leq(Ident("x"), Const(0.0f)) // x <= 0.0
val stmt3 = AssignStmt("x", Minus(Const(0.0f), Ident("x"))) // x= 0 - x
val ifStmt = IfThenElseStmt(cond2, List(stmt3), List()) // if ( x <= 0) x = -x else nothing
val whileStmt = WhileStmt(cond1, List(stmt2, ifStmt)) // while loop 
val allstmts = List(stmt1, whileStmt)
val retExpr = Minus(Ident("y"), Ident("x")) // y - x to be returned

val c1 = SwitchCase(List(Const(2.0), Const(3.0)), List(stmt1, stmt3))
val c3 = SwitchCase(List(Ident("y"), Ident("z")), List(stmt1, stmt2, whileStmt))
val c2 = DefaultCase(List(ifStmt))

// Check if we can write a multiway switch
val v1 = Switch(List(e1,e2), List(c1, c2, c3))

passed(5)
//END TEST
```

Output:
```
*** Tests Passed (5 points) ***
[36me1[39m: [32mPlus[39m = [33mPlus[39m(
  [33mIdent[39m([32m"y"[39m),
  [33mPlus[39m([33mIdent[39m([32m"z"[39m), [33mExp[39m([33mMinus[39m([33mIdent[39m([32m"x"[39m), [33mIdent[39m([32m"y"[39m))))
)
[36me2[39m: [32mDiv[39m = [33mDiv[39m([33mPlus[39m([33mIdent[39m([32m"x"[39m), [33mIdent[39m([32m"y"[39m)), [33mConst[39m([32m2.0[39m))
[36mstmt1[39m: [32mAssignStmt[39m = [33mAssignStmt[39m(
  [32m"x"[39m,
  [33mPlus[39m([33mIdent[39m([32m"y"[39m), [33mPlus[39m([33mIdent[39m([32m"z"[39m), [33mExp[39m([33mMinus[39m([33mIdent[39m([32m"x"[39m), [33mIdent[39m([32m"y"[39m)))))
)
[36mcond1[39m: [32mLeq[39m = [33mLeq[39m([33mIdent[39m([32m"y"[39m), [33mConst[39m([32m15.0[39m))
[36mstmt2[39m: [32mAssignStmt[39m = [33mAssignStmt[39m([32m"y"[39m, [33mMinus[39m([33mIdent[39m([32m"y"[39m), [33mIdent[39m([32m"x"[39m)))
[36mcond2[39m: [32mLeq[39m = [33mLeq
```

## Problem 2 (15 Points)

Add list manipulation expressions to the allowed expressions in the `mython` language. 

We allow list expressions (use a new non terminal called __ListExpr__).

Formally, a list expression is defined as one of the following:
- The **empty list**: use a constructor symbol `EmptyList`. It should not have arguments.
- A  **new list** made up of zero or more arithmetic expressions: use constructor symbol `NewList`. It has as arguments a list of arithmetic expressions. Use Kleene-star in your grammar to specify "one or more of".
- Appending an arithmetic expression to the front of a list: use constructor symbol `Cons`. `Cons` must have two arguments: first argument is an arithmetic expression and second argument is a list expression.
- Concatenate two lists defined by list expressions: use constructor symbol `Concat`. It must have two arguments that are both list expressions.
- Filter elements of a list which meet a certain condition: use constructor symbol `Filter`. It must have three arguments: (a) a list expression denoting the list being filtered, (b) an identifier name (use the nonterminal __Identifier__), and (c) a conditional expression capturing the filtering criterion. We would use such an expression, for instance, as following: `filter (x in my_list) where x>0`. This would create a new list with all elements `x` from `my_list` where `x>0`.

Next, we add the following rule to  __CondExpr__  (conditional expressions):
- A condition expression to check if a list is empty:  Use constructor `IsEmptyList`. Its argument must be a list expression that we check for emptiness.

Finally, we add a special foreach loop statement over the elements of a list.

`foreach (x in list_expr) begin  .. list of statements...  end `

- Add a statement to loop over the elements of a list: use constructor`ForEach`. Its arguments must include (a) the identifier, (b) the list expression to be iterated over and (c) the body of the loop (list of statements).

Write down the extra grammar rules that you would add to the overall 
`mython` grammar to create these functionalities.

<!-- grade_id: cell-76cfd1d2c6a5898d, points: 20, grade: true, solution: true -->
YOUR ANSWER HERE

$$\begin{array}{rcl}
\mathbf{ListExpr} & \rightarrow & \text{EmptyList} \\
& | & \text{NewList} (\mathbf{Expr}^*)\\
& | & \text{Cons} (\mathbf{Expr},\mathbf{ListExpr})\\
& | & \text{Concat} (\mathbf{ListExpr} , \mathbf{ListExpr})\\
& | & \text{Filter} (\mathbf{ListExpr},\mathbf{Identifier},\mathbf{CondExpr})\\
\mathbf{CondExpr} & \rightarrow & \text{IsEmptyList}(\mathbf{ListExpr}) \\
\mathbf{Statement} & \rightarrow & \text{ForEach}(\mathbf{Identifier},\mathbf{ListExpr},\mathbf{Statement}^*) \\
\end{array}$$

## Problem 3 (20 points)

### A (10 points)

We defined lists in the class as an inductive datatype.  

Let's define a  `NumOption` type for functions with optional return values (such as `Some(2)` or `None`).  See note on option types as the inspiration for NumOption: https://www.geeksforgeeks.org/scala-option/

We will use the following grammar:

$$\begin{array}{rcll}
\mathbf{NumOption} & \Rightarrow & \textit{None} & \text{denotes the lack of a value }\\
& |& \textit{Some}(\mathbf{Integer}) & \text{denotes an integer value } \\
\end{array} $$

First define, the `NumOption` datatype following the grammar above.


Next write a recursive function that returns the nth element of the list. However, if the list size is smaller than `n` you should return a `None` signalling no element can be found or otherwise `Some(x)` where `x` is the nth element. Also, if `n < 0`,  throw an `IllegalArgumentException` conveying the appropriate endearments to the user in the form of a string message. 

Assume $n=0$ denotes the very first element and $n = \text{length of list} -1$ denotes the very last element. 

<!-- grade_id: cell-724b4ca89445ceaf, points: 0, grade: false, solution: true -->
```scala
sealed trait NumList
case object Nil extends NumList
case class Cons(n: Int, l: NumList) extends NumList 

/*-- Define NumOption data type first --*/
// YOUR CODE HERE
sealed trait NumOption
case object None extends NumOption
case class Some(n:Int) extends NumOption

/*-- Next, implement this function: see problem description for specification --*/
def getNthElement(lst: NumList, n: Int): NumOption = {
    // YOUR CODE HERE
    //if(lst.size < n){None}
    if(n<0) throw new IllegalArgumentException
    lst match
    {
         case Nil => None
        
       case Cons(l1,lists) if(n==0)=>  Some(l1)
        case Cons(l1,lists) if(n>=1)=> {getNthElement(lists,n-1)}
    }
    
}
```

Output:
```
defined [32mtrait[39m [36mNumList[39m
defined [32mobject[39m [36mNil[39m
defined [32mclass[39m [36mCons[39m
defined [32mtrait[39m [36mNumOption[39m
defined [32mobject[39m [36mNone[39m
defined [32mclass[39m [36mSome[39m
defined [32mfunction[39m [36mgetNthElement[39m
```

<!-- grade_id: cell-2bc6d5ee94793b48, points: 10, grade: true, solution: false -->
```scala
//BEGIN TEST
val l1 = Nil
val l2 = Cons(1, Cons(-1, Nil))
val l3 = Cons(1, Cons(2, l2))
val l4 = Cons(0, Cons(4, Cons(8, l3)))

val test1 = try {
    getNthElement(l2, -1);
    assert(false, "Test 1 : getNthElement(l2, -1) should raise an IllegalArgumentException. Your code did not.")
} catch {
    case e: IllegalArgumentException => "OK"
} 
assert(getNthElement(l1, 3) == None, "Test2: getNthElement(l1, 3)  failed (expected answer = None)")
assert(getNthElement(l2, 0) == Some(1), "Test3: getNthElement(l2, 0)  failed (expected answer = Some(1))")
assert(getNthElement(l3, 3) == Some(-1), "Test4: getNthElement(l3, 3)  failed (expected answer = Some(-1))")
assert(getNthElement(l4, 2) == Some(8), "Test5: getNthElement(l4, 2)  failed (expected answer = Some(8))")

val test6 = try {
    getNthElement(l4, -8);
    assert(false, "Test 6 : getNthElement(l4, -8) should raise an IllegalArgumentException. Your code did not.")
} catch {
    case e: IllegalArgumentException => "OK"
}

passed(5)
//END TEST
```

Output:
```
*** Tests Passed (5 points) ***
[36ml1[39m: [32mNil[39m.type = Nil
[36ml2[39m: [32mCons[39m = [33mCons[39m([32m1[39m, [33mCons[39m([32m-1[39m, Nil))
[36ml3[39m: [32mCons[39m = [33mCons[39m([32m1[39m, [33mCons[39m([32m2[39m, [33mCons[39m([32m1[39m, [33mCons[39m([32m-1[39m, Nil))))
[36ml4[39m: [32mCons[39m = [33mCons[39m([32m0[39m, [33mCons[39m([32m4[39m, [33mCons[39m([32m8[39m, [33mCons[39m([32m1[39m, [33mCons[39m([32m2[39m, [33mCons[39m([32m1[39m, [33mCons[39m([32m-1[39m, Nil)))))))
[36mtest1[39m: [32mAny[39m = [32m"OK"[39m
[36mtest6[39m: [32mAny[39m = [32m"OK"[39m
```

### B (7 points)
Write a recursive procedure that returns true if and only if the list has consecutive duplicates, i.e., a number occurs at least twice in succession. 

Note that the property is trivially "false" for lists of sizes $0$ and $1$.

#### Example

  - `List(0, 1, 1, 2, 3, 4, 5, 2)` has consecutive duplicates "1" followed by "1".
  - `List(0, 1, 0, 1, 0, 1, 0, 1)` does not have consecutive duplicates (although it has duplicates).



<!-- grade_id: cell-eedd55e1f86bd6c6, points: 0, grade: false, solution: true -->
```scala
def hasConsecutiveDupes(lst: NumList): Boolean = {
    // YOUR CODE HERE
   lst match
    {
        case Nil => return false
        case Cons(l1,Nil) => return false
        case Cons(l1,Cons(l2,Nil)) => return false
        case Cons(l1,Cons(l2,Cons(l3,lists))) if(l1==l2)=> return true
        case Cons(l1,Cons(l2,Cons(l3,lists))) if(l1!=l2)=> hasConsecutiveDupes(Cons(l2,Cons(l3,lists)))
    }
}
```

Output:
```
defined [32mfunction[39m [36mhasConsecutiveDupes[39m
```

<!-- grade_id: cell-33a4f21c01f0c7c4, points: 7, grade: true, solution: false -->
```scala
//BEGIN TEST
val l1 = Cons(12, Cons(25, Cons(12, Nil)))
assert(!hasConsecutiveDupes(l1), 
       "Test case 1 :  hasConsecutuveDupes(l1) -- should return false")

val l2 = Cons(14, Cons(14, Cons(13, l1 )))
assert(hasConsecutiveDupes(l2), 
       "Test case 2 :  hasConsecutiveDupes(l2) -- should return true")

val l3 = Cons(0, Cons(-2, Cons(-2, l1 )))
assert(hasConsecutiveDupes(l3), 
       "Test case 3 :  hasConsecutiveDupes(l2) -- should return true")

assert(!hasConsecutiveDupes(Nil), 
       "Test case 4 :  hasConsecutiveDupes(Nil) -- should return false")
val l5 = Cons(-1, Nil)
assert(!hasConsecutiveDupes(l5), 
       "Test case 5:  hasConsecutiveDupes(l5) -- should return false")

passed(7)
//END TEST
```

Output:
```
*** Tests Passed (7 points) ***
[36ml1[39m: [32mCons[39m = [33mCons[39m([32m12[39m, [33mCons[39m([32m25[39m, [33mCons[39m([32m12[39m, Nil)))
[36ml2[39m: [32mCons[39m = [33mCons[39m([32m14[39m, [33mCons[39m([32m14[39m, [33mCons[39m([32m13[39m, [33mCons[39m([32m12[39m, [33mCons[39m([32m25[39m, [33mCons[39m([32m12[39m, Nil))))))
[36ml3[39m: [32mCons[39m = [33mCons[39m([32m0[39m, [33mCons[39m([32m-2[39m, [33mCons[39m([32m-2[39m, [33mCons[39m([32m12[39m, [33mCons[39m([32m25[39m, [33mCons[39m([32m12[39m, Nil))))))
[36ml5[39m: [32mCons[39m = [33mCons[39m([32m-1[39m, Nil)
```

### C (8 points)
Write a recursive function `filterNumList(l: NumList, f: Int => Boolean): NumList` that takes in a `NumList` and a function `f: Int => Boolean`. 

1. It should return a new list that consist of all elements of the list `l` that return `true` when the function `f` is called on them.
2. The returned list elements must preserve the same order as in the original list.

<!-- grade_id: cell-0f9f5a51ae34b3bb, points: 0, grade: false, solution: true -->
```scala
// YOUR CODE HERE
def filterNumList(l: NumList, f: Int => Boolean): NumList=
{
    l match
    {
        case Nil => Nil //empty
        case Cons(l1,lists) if(f(l1))=>Cons(l1,filterNumList(lists,f))//l1 is true,then call cons function, put l1 in
        case Cons(l1,lists) if(!f(l1)) => filterNumList(lists,f)// l1 is false,recurive the function
    }
}
```

Output:
```
defined [32mfunction[39m [36mfilterNumList[39m
```

<!-- grade_id: cell-a2549ef77f3f9bb0, points: 8, grade: true, solution: false -->
```scala
//BEGIN TEST
val l1 = Cons(12, Cons(25, Cons(37, Nil)))
def f1(j: Int): Boolean =  j <= 25 && j >= 12
assert(filterNumList(l1, f1) == Cons(12, Cons(25, Nil)), "Test 1 failed.")


val l2 = Cons(22, Cons(135, Cons(137, l1)))
def f2(j: Int): Boolean =  j % 5 == 0
assert(filterNumList(l2, f2) == Cons(135, Cons(25, Nil)), "Test 2 failed.")

def f3(j: Int): Boolean =  j >= 210
assert(filterNumList(l2, f3) == Nil, "Test 3 failed.")
assert(filterNumList(Nil, f3) == Nil, "Test 4 failed.")

val l4 = Cons(0, Cons(0, Cons(0, Cons(0, Cons(0, Cons(0, Nil))))))

def f4(j: Int): Boolean =  j <= 0
assert(filterNumList(l4, f4) == l4, "Test 5 failed")

passed(8)
//END TEST
```

Output:
```
*** Tests Passed (8 points) ***
[36ml1[39m: [32mCons[39m = [33mCons[39m([32m12[39m, [33mCons[39m([32m25[39m, [33mCons[39m([32m37[39m, Nil)))
defined [32mfunction[39m [36mf1[39m
[36ml2[39m: [32mCons[39m = [33mCons[39m([32m22[39m, [33mCons[39m([32m135[39m, [33mCons[39m([32m137[39m, [33mCons[39m([32m12[39m, [33mCons[39m([32m25[39m, [33mCons[39m([32m37[39m, Nil))))))
defined [32mfunction[39m [36mf2[39m
defined [32mfunction[39m [36mf3[39m
[36ml4[39m: [32mCons[39m = [33mCons[39m([32m0[39m, [33mCons[39m([32m0[39m, [33mCons[39m([32m0[39m, [33mCons[39m([32m0[39m, [33mCons[39m([32m0[39m, [33mCons[39m([32m0[39m, Nil))))))
defined [32mfunction[39m [36mf4[39m
```

### That's All Folks !!!