Before you turn this problem in, make sure everything runs as expected. 
  1. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and 
  2. Then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name below. 

---

# CSCI 3155 Assignment 1 : Scala Programming

This assignment asks you to write scala programs. Fill in the programs below. Use the test cases provided to test them. You are also encouraged to write your own test cases.

**Important** It is important to run the cells in order. In particular, run the cell that defines the functions passed and testWithMessage first. If not, you will get errors saying these functions are not defined.

**YOUR NAME HERE**

```scala
// TEST HELPER
// IMPORTANT: Run this cell first before attempting anything else. 
// Or else, you will get error messages saying passed: not found etc..
def passed(points: Int) {
    require(points >=0)
    if (points == 1) print(s"\n*** Tests Passed (1 point) ***\n")
    else print(s"\n*** Tests Passed ($points points) ***\n")
}

def testWithMessage(cond: Boolean, testID: String) = { 
    assert(cond, s"Test $testID Failed"); 
    println(s"Test $testID Passed!")
}
```

Output:
```
defined [32mfunction[39m [36mpassed[39m
defined [32mfunction[39m [36mtestWithMessage[39m
```

## Problem 1 (10 points)

Write a scala function that searches for an occurrence of the substring `pat` inside the string `s`. It returns a true if the pattern `pat` is a substring of `s` and false otherwise.

Useful functions on strings that are permitted.
  - str.size : number of characters in a string `str`.
  - str(i): the char at ith index of str. i between 0 and str.size -1
  
You can use a nested for/while loop with vars and return statements for this problem.
Note about for loops in scala:
  - `for (i <- 0 until n) { ...  } ` starts from 0 and runs to n-1.
  -  `for (i <- 0 to n) { ...  } ` starts from 0 and runs to n.
  


**You are forbidden from using string library inbuilt functions ( other than the ones mentioned above ) such as substring that will directly solve the problem, or other tricks such as converting a string to a list and using list API functions etc.. Use of for loops and basic indexing of strings is permitted.**


<!-- grade_id: cell-794743da17722228, points: 0, grade: false, solution: true -->
```scala
def patternOccursInString(s: String, pat: String): Boolean = {
    // YOUR CODE HERE
    val n:Int= s.size;
    val m:Int=pat.size;
  
    if(m>n){return false}
    else if(pat==""){return true}
    else{
    for (  i <- 0 until n-m+1)
    {
         var y:Boolean=true
         for(j <-0 until m)
        {
            if(s(i+j)!=pat(j))
            {

               y=false
                
            }
            
        }
       if(y==true){return true}
       
      }
      return false
    }
}  
```

Output:
```
defined [32mfunction[39m [36mpatternOccursInString[39m
```

<!-- grade_id: cell-082571ccc4ca97d0, points: 10, grade: true, solution: false -->
```scala
//BEGIN TEST
assert(patternOccursInString("bahbahblacksheep", "hbla"), "Test 1 Failed")
assert(patternOccursInString("wiajsioioia", "j"), "Test 2 Failed")
assert(!patternOccursInString("woia", "woiaaaalapo"), "Test 3 Failed")
assert(patternOccursInString("papajohnspizza",""), "Test 4 Failed")
assert(patternOccursInString("goatwearstoga", "toga"), "Test 5 Failed")
assert(!patternOccursInString("aaaaaaax", "axe"), "Test 6 Failed")
assert(patternOccursInString("",""), "Test 7 failed")
passed(10)
//END TEST
```

Output:
```
*** Tests Passed (10 points) ***
```

```scala

```

## Problem 2 (20 points)

Implement a `patternOccursInStringRec(s, pat)` function that checks if a given pattern `pat` occurs in the string `s`, but 
  - without using mutable `var`, 
  - without using `for`/`while` loops, and
  - without using the `return` keyword.

Your recursive function should use the following high level strategy:
 
First write a helper function `searchPatternAtPosRec(s, pat, i)` that searches if pattern occurs starting at position `i` of string `s`.
  - if pat is the empty string return true
  - else if i >= s.size and pat is not empty string, return false.
  - else, if s(i) != pat(0) return false
  - else, recursively call searchPatternAtPosRec(s, pat.substring(1), i+1)
    - Note `pat.substring(1)` can be used and takes the substring of `pat` from index 1 until the end.
 
Next, implement the recursive function `patternOccursInStringRec(s, pat,j)`
  - if pat is the empty string, then return true
  - else if j >= s.size then return false
  - else if (searchPatternAtPosRec(s, pat, j) is true) return true
  - else recursively call patternOccursInStringRec by incrementing j
 

<!-- grade_id: cell-80da0f71ccafab19, points: 0, grade: false, solution: true -->
```scala
// YOUR CODE HERE
def searchPatternAtPosRec(s:String,pat:String,i:Int):Boolean=
{
    val n:Int=s.length
    if(pat==""){true}
    else if(i>= n){false}
    else if(s(i)!=pat(0)){false}
    else searchPatternAtPosRec(s,pat.substring(1),i+1)
    
}
```

Output:
```
defined [32mfunction[39m [36msearchPatternAtPosRec[39m
```

<!-- grade_id: cell-ef781cc6d55287ba, points: 10, grade: true, solution: false -->
```scala
//BEGIN TEST

testWithMessage(searchPatternAtPosRec("hello", "ll", 2), "1")
testWithMessage(!searchPatternAtPosRec("hello", "ll", 4), "2")
testWithMessage(!searchPatternAtPosRec("hello", "oo", 2), "3")
testWithMessage(searchPatternAtPosRec("hello", "hello", 0), "4")
testWithMessage(searchPatternAtPosRec("hello", "hell", 0), "5")
testWithMessage(searchPatternAtPosRec("hello", "", 4), "6")
passed(10)
//END TEST
```

Output:
```
Test 1 Passed!
Test 2 Passed!
Test 3 Passed!
Test 4 Passed!
Test 5 Passed!
Test 6 Passed!

*** Tests Passed (10 points) ***
```

<!-- grade_id: cell-79fc1ba65f9c2f3e, points: 0, grade: false, solution: true -->
```scala
def patternOccursInStringRec(s: String, pat: String, j : Int = 0): Boolean = {
    // YOUR CODE HERE
    val n:Int=s.length
    if(pat==""){true}
    else if(j>=n ){false}
    else if (searchPatternAtPosRec(s, pat, j)==true ) {true} 
    else patternOccursInStringRec (s,pat,j+1)
}
```

Output:
```
defined [32mfunction[39m [36mpatternOccursInStringRec[39m
```

<!-- grade_id: cell-3a1b07140e65fe6c, points: 10, grade: true, solution: false -->
```scala
assert(patternOccursInStringRec("bahbahblacksheep", "hbla"), "Test 1 Failed")
assert(patternOccursInStringRec("wiajsioioia", "j"), "Test 2 Failed")
assert(!patternOccursInStringRec("woia", "woiaaaalapo"), "Test 3 Failed")
assert(patternOccursInStringRec("papajohnspizza",""), "Test 4 Failed")
assert(patternOccursInStringRec("goatwearstoga", "toga"), "Test 5 Failed")
assert(!patternOccursInStringRec("aaaaaaax", "axe"), "Test 6 Failed")
assert(patternOccursInStringRec("",""), "Test 7 failed")
passed(10)
```

Output:
```
*** Tests Passed (10 points) ***
```

## Problem 3 (20 points)

A polynomial of the form $a_0 + a_1 x + a_2 x^2 + ... + a_n x^n$ is 
represented by a list of coefficients `[a0, ..., an]`. We wish to implement a polynomial class in Scala. Complete the missing methods.

There are no restrictions in the use of loops/var in this problem. Also any list API function may be used.

**Note:** We will guarantee in our implementation that the highest coefficient `an != 0`. This will simplify things a lot.

### (A, 5 points) Chop off suffix zeros

First implement a function `chopOffSuffixZeros(lst: List[Double]): List[Double]` that given a list of numbers, removes elements from the end that are equal to 0.0, ensuring that the returned list has a non zero element at its last position.

Examples: 

`chopOffSuffixZeros( List(1.0, 2.1, -1.0, 0.0, 0.0) ) ` should return 
`List(1.0, 2.1, -1.0)`.

`chopOffSuffixZeros( List(0.0, 0.1, 0.0, 0.0, 0.0) ) ` should return 
`List(0.0, 0.1)`.

`chopOffSuffixZeros( List(0.0, 0.0, 0.0, 0.0, 0.0, 1.0) ) ` should return 
`List(0.0, 0.0, 0.0, 0.0, 0.0, 1.0)`.

`chopOffSuffixZeros( Nil ) ` should return `Nil` (Nil stands for empty list in scala).


- You may use for /while loops, var etc..
- You can use list API functions such as `slice`, `last`, `size`, list concatenation, etc..
- However, this is an example of a problem where the recursive solution is way easier than the one using for loop. Try it out!!



<!-- grade_id: cell-bbd6940ce326dfbb, points: 0, grade: false, solution: true -->
```scala
def chopOffSuffixZeros(lst: List[Double]): List[Double] = {
    // YOUR CODE HERE
    val n:Int=lst.length;
    if(n==0){return Nil}
    else if(lst.last != 0.0){return lst}
   else chopOffSuffixZeros(lst.slice(0,n-1))
} 
```

Output:
```
defined [32mfunction[39m [36mchopOffSuffixZeros[39m
```

<!-- grade_id: cell-234d164ebb0350f1, points: 5, grade: true, solution: false -->
```scala
testWithMessage(
    chopOffSuffixZeros( List(1.0, 2.1, -1.0, 0.0, 0.0) ) 
                       == List(1.0, 2.1, -1.0), "1")


testWithMessage(
    chopOffSuffixZeros(   List(1.0, 2.1, -1.0) ) 
                        == List(1.0, 2.1, -1.0), "2")

testWithMessage(
    chopOffSuffixZeros(   List(0.0, 0.0, 0.0, 0.0, 0.0, 0.0, -0.1) ) 
                       == List(0.0, 0.0, 0.0, 0.0, 0.0, 0.0, -0.1), "3")


testWithMessage(
    chopOffSuffixZeros( List() ) == List(), "4")

testWithMessage(
    chopOffSuffixZeros( List(0.0, 0.0, 0.0, 0.0, 0.0, 0.0) ) == List(), "5")

passed(5)

```

Output:
```
Test 1 Passed!
Test 2 Passed!
Test 3 Passed!
Test 4 Passed!
Test 5 Passed!

*** Tests Passed (5 points) ***
```

### Part(B)

Implement the missing methods in the polynomial class below.

<!-- grade_id: cell-a410d01fa666b077, points: 0, grade: false, solution: true -->
```scala
class Polynomial(val coeffs: List[Double]) {
    // Implement a pretty printer for the class
    // by overriding the toString() method
    override def toString(): String = {
        coeffs.zipWithIndex
            .map { case (aj, j) => s"($aj* x^$j)"}
            .mkString("+")
    }
    
    def polyEquals(poly: Polynomial): Boolean = {
        //Check if two polynomials are equal by comparing the list of 
        // coefficients.
        // Assume that both this and poly have no suffix zeros to make it simple.
        // YOUR CODE HERE
        if(coeffs== (poly.coeffs)) return true
        else return false
        
    }
    
    // This is an operator overloading for the + operator.
    def + (poly2: Polynomial): Polynomial = {
        //TODO: Return a new polynomial that is the sum of 
        //   the coefficients of this polynomial and those of poly2
        // Use a for/while loop but make sure to handle the case
        //        when the sizes of poly and poly2 are different.
        // Finally, make sure that you chop off all zeros at the
        // end using chopOffSuffixZero funcion.
        // YOUR CODE HERE
        var n:Int=poly2.coeffs.size;
        var m:Int= coeffs.size;
     var temp:List[Double]=Nil;
      /*for(i<-0 until m) 
        {
            var temp = coeffs(i) ++ poly2.coeffs(i) ++list
            list=temp
        }*/
        if(n==m)
        {
            for(i<- 0 until n)
            {
                temp = temp ++  List(coeffs(i) + poly2.coeffs(i))
               
            }
            
        }
       
       else if(m<n)
        {
       
           for(i<- 0 until m)
            {
                temp = temp ++  List(coeffs(i) + poly2.coeffs(i))
               
            }
         
                for(j<- m until n)
            {
                  temp=temp  ++ List(poly2.coeffs(j).toDouble)
                
            }
              
            
            
        }
        else
        {  
            for(i<- 0 until n)
            {
                temp = temp ++  List(coeffs(i) + poly2.coeffs(i))
               
            }
      
              for(j<- n until m)
            {
                 temp =temp  ++ List(coeffs(j) .toDouble)
                
            }
            
           
        }
   def chopOffSuffixZeros(lst: List[Double]): List[Double] = {
    // YOUR CODE HERE
    val n:Int=lst.length;
    if(n==0){return Nil}
    else if(lst.last != 0.0){return lst}
   else chopOffSuffixZeros(lst.slice(0,n-1))
} 
        new Polynomial(chopOffSuffixZeros(temp))
        
    }
   
    // We can use your rationalEquals method to overide the == operator as follows
    // Ignore this method for now -- it will be clearer soon
    override def equals(r: Any): Boolean = r match {
        case poly: Polynomial => this.polyEquals(poly) 
        case _ => false
    }
    // END Ignore   
}
```

Output:
```
defined [32mclass[39m [36mPolynomial[39m
```

<!-- grade_id: cell-6275b75f19c16246, points: 5, grade: true, solution: false -->
```scala
//BEGIN TEST
val p = new Polynomial(List(0, 1, -2))
val p2 = new Polynomial(List(0, 1, -1))
val p3 = new Polynomial(List(0, 1, -2))
testWithMessage(p != p2, "1")
testWithMessage(p == p3, "2")
passed(5)
//END TEST
```

Output:
```
Test 1 Passed!
Test 2 Passed!

*** Tests Passed (5 points) ***
[36mp[39m: [32mPolynomial[39m = (0.0* x^0)+(1.0* x^1)+(-2.0* x^2)
[36mp2[39m: [32mPolynomial[39m = (0.0* x^0)+(1.0* x^1)+(-1.0* x^2)
[36mp3[39m: [32mPolynomial[39m = (0.0* x^0)+(1.0* x^1)+(-2.0* x^2)
```

<!-- grade_id: cell-f595d18030ee424d, points: 10, grade: true, solution: false -->
```scala
//BEGIN TEST
val p1 = new Polynomial(List(0, 1, 2))
val p2 = new Polynomial(List(0, 1, -1))
val p3 = new Polynomial(List(0, 1, -2, 3))
val p4 = p1 + p2
testWithMessage(p4.coeffs == List(0.0, 2.0, 1.0), "1")
val p5 = p2 + p3
testWithMessage(p5.coeffs == List(0.0, 2.0, -3.0, 3.0), "2")
val p6 = new Polynomial(List(-1.0, -2.0, 3.0, -3.0))
val p7 = p5 + p6
testWithMessage(p7.coeffs == List(-1.0), "3")
val p8 = new Polynomial(List(1.0))
val p9 = p7 + p8
testWithMessage(p9.coeffs == Nil, "4")
passed(10)
//END TEST
```

Output:
```
Test 1 Passed!
Test 2 Passed!
Test 3 Passed!
Test 4 Passed!

*** Tests Passed (10 points) ***
[36mp1[39m: [32mPolynomial[39m = (0.0* x^0)+(1.0* x^1)+(2.0* x^2)
[36mp2[39m: [32mPolynomial[39m = (0.0* x^0)+(1.0* x^1)+(-1.0* x^2)
[36mp3[39m: [32mPolynomial[39m = (0.0* x^0)+(1.0* x^1)+(-2.0* x^2)+(3.0* x^3)
[36mp4[39m: [32mPolynomial[39m = (0.0* x^0)+(2.0* x^1)+(1.0* x^2)
[36mp5[39m: [32mPolynomial[39m = (0.0* x^0)+(2.0* x^1)+(-3.0* x^2)+(3.0* x^3)
[36mp6[39m: [32mPolynomial[39m = (-1.0* x^0)+(-2.0* x^1)+(3.0* x^2)+(-3.0* x^3)
[36mp7[39m: [32mPolynomial[39m = (-1.0* x^0)
[36mp8[39m: [32mPolynomial[39m = (1.0* x^0)
[36mp9[39m: [32mPolynomial[39m =
```

-- That is all for now, folks! -- 

```scala

```